Mid-zone hepatocytes trade proliferation for survival via Atf4-Chop axis in early acute liver injury

Author response:

The following is the authors’ response to the original reviews.

Public Reviews:

Reviewer #1 (Public review):

(1) Zonation definition under injury has been shown to be sustained broadly, but is not sufficiently validated and quantified, especially considering the resolution of the 10x Visium system and the potential variation of outcomes based on how to define zones.

We thank the reviewer for these insightful suggestions. In this study, under normal conditions (APAP 0h), each liver lobule was divided into three zones based on unbiased gene expression profiles. The PP zone was defined by enrichment of PP signature genes (e.g.,Alb, Mup20, Cyp2f2,Pck1,Apoa4). The PC zone was defined by high expression of PC markers (e.g., Gs, Cyp2e1, Oat, Cyp1a2, Apoe). The Mid zone comprised regions with intermediate expression of PC and PP markers and elevated levels of Igfbp2 and Hamp (Revised Figure 2A and S1C). Following APAP-induced injury (3 h, 6 h), the PC zone remained identifiable based on residual enrichment of PC signature genes (e.g., Cyp1a2,Glul) despite necrosis and reduced overall transcription, while PP gene expression remained largely unchanged. The Mid zone was defined as the transcriptional cluster between PC and PP regions exhibiting marked reprogramming, (e.g.,Sqstm1, Igfbp1) (Revised Figure 2A and S1C). To validate and quantify our zonation approach, we compared it with classical nine even layers from central vein (CV) to portal vein (PV). Immunostaining and quantification for Cyp2f2 (a PP marker), p62 (the protein product of Sqstm1, a Mid marker during early liver injury), Glutamine Synthetase (GS, the protein product of Glul, a PC marker) further corroborated zone definitions at each time point, showing correspondence of our PC (layers 1–2), Mid (layers 3–6), and PP (layers 7–9) (Revised Figure 2 B-D) (Revised manuscript, page 5, lines 119–131, page 6-7, line 174-182).

(2) The model is built entirely in APAP injury, which specifically targets pericentral hepatocytes. It remains unclear whether the proposed mechanism applies to other liver injuries (e.g., partial hepatectomy, CCl4).

We thank the reviewer for this insightful comment. To test whether the proposed mechanism applies to other liver injuries, we employed mouse models of partial hepatectomy (PHx) and carbon tetrachloride (CCl4)-induced acute liver injury. In our CCl4 model (administered intraperitoneally in corn oil, with samples collected 18 h post‑injection), the ISR was activated around injury sites, accompanied by decreased proliferation, as evidenced by increased expression of p‑eIF2α, Atf4, Chop, and Btg2, along with reduced Ki67 expression (Revised Figure S6A–G). In PHx model (examined 24 h after surgery), ISR activation was similarly observed around ischemic injury sites, with increased p‑eIF2α, Atf4, Chop, and Btg2 expression and undetectable Ki67 expression (Revised Figure S7A–G). Together, these additional models suggest that the proposed mechanism may be applicable to other types of liver injury (Revised manuscript, page 15, Line 410-426).

(3) Baseline proliferation appears higher than expected in homeostasis (Figure 1B), and fold change analysis (not absolute counts) may be needed to assess zonal proliferation suppression (Figure 1D).

We thank the reviewer for this insightful comment. The baseline proliferation rates observed in our study are consistent with previously reported zonal distributions (PMID: 33632817; PMID: 33632818), with approximately 70% of proliferating hepatocytes located in zone 2, 20% in zone 3, and 10% in zone 1 under homeostatic conditions. To further address the reviewer’s concern, we performed a fold-change analysis of Ki-67+hepatocytes across different zones. This analysis revealed that only the mid (zone 2) and pericentral regions exhibited significant changes, whereas no statistically significant differences were observed in the other zones (as shown in Author response image 1). Importantly, when considered together with the absolute cell counts, these results indicate that the apparent suppression of proliferation is most pronounced in the mid zone, likely due to its relatively higher baseline proliferation under homeostatic conditions. In contrast, this effect is less evident in the fold-change analysis, as zones with low baseline proliferation show limited dynamic range for detecting relative changes.

Fold changes of Ki67-positive cells across liver zones (PC, Mid, PP) at 0, 3, 6, 12 and 24 h post-APAP.

Fold change was the number of Ki67-positive cells in the three regions at each time point after APAP treatment divided by the number of positive cells in each region at 0 hour post-APAP. (a) denotes significance between PC and Mid regions, (b) denotes significance between PC and PP regions, and (c) denotes significance between Mid and PP regions.

(4) AAV-based overexpression raises potential confounds (altered CYP activity before injury) and shows incomplete penetrance that is not quantified (Figure 5 - Figure 6).

We thank the reviewer for raising these important points. We measured basal Cyp2e1 protein levels by western blot in AAV‑EGFP, AAV‑Atf4, and AAV‑Btg2 mice without APAP treatment. Compared to AAV‑EGFP controls, Cyp2e1 expression was modestly reduced in the Atf4 and Btg2 groups, respectively (Revised Figure S5A). Although we assessed protein abundance rather than enzymatic activity directly, Cyp2e1 protein levels under basal conditions generally correlate well with activity. Published studies demonstrate that robust protection against APAP hepatotoxicity typically requires >50% suppression of CYP2E1 activity (PMID: 35145060; PMID: 30151903). The minor reductions we observed are therefore far below the threshold needed to explain the 70–90% decreases in serum ALT conferred by Atf4 or Btg2 overexpression (Revised Figures 5D and 6I). Accordingly, altered CYP2E1 activity is unlikely to represent a significant confound in our model.

We quantified transduction efficiency by immunohistochemical detection of the respective transgene proteins and determined the percentage of positive hepatocytes. At a dose of 1.2 × 1011 viral genomes per animal, average transduction rates were 32% (EGFP), 18% (Atf4), and 23% (Btg2) (Revised Figure S5B). Individual animal transduction efficiency showed a negative correlation with serum ALT levels (e.g. Pearson r = –0.7681, p = 0.0260 for Atf4; Revised Figure S5C), demonstrating that greater transgene expression associates with stronger protection. Although these average transduction rates appear modest relative to the 70–90% reduction in ALT, this apparent disproportion is consistent with the known tendency of AAV‑TBG vectors to transduce hepatocytes preferentially in the pericentral region—the same zone where APAP‑induced necrosis initiates. Pericentral enrichment of transgene expression could thus provide disproportionate protection by targeting the most vulnerable cells. These data are now included in Revised Figure S5A–C and detailed in the Results (page 14, lines 383–399).

(5) The functional link between proliferation suppression and improved survival is inferred, but direct survival /injury readouts are limited.

We thank the reviewer for this insightful comment. To more directly evaluate the functional link between proliferation control and liver injury, we manipulated Btg2, a downstream effector of the Atf4–Chop axis and a known inhibitor of cell proliferation. Knockdown of Btg2 using AAV8–CasRx achieved a moderate (~30%) reduction in Btg2 expression (Revised Figure S5D). Despite this partial knockdown efficiency, we observed a clear exacerbation of liver injury, as evidenced by an approximately 2-fold increase in serum ALT levels and a ~1.5-fold expansion of necrotic areas. In parallel, hepatocyte proliferation was significantly increased (~1.8-fold increase in Ki67⁺ hepatocytes) compared to control mice (Revised Figure 6K–N). Conversely, Btg2 overexpression produced the opposite phenotype, markedly attenuating liver injury while suppressing hepatocyte proliferation (Revised Figure 6G–J). Together, these gain- and loss-of-function data provide direct evidence linking proliferation control to injury severity, thereby supporting a causal relationship between suppressed proliferation and improved liver outcomes (Revised manuscript, page 14, lines 399–406).

Reviewer #2 (Public Review):

(1) Starting with the basics, one wonders why midlobular hepatocytes manage to mount a defensive response to APAP but pericentral hepatocytes don't. Is this because midlobular hepatocytes express the relevant Cyps (2e1, but also 1a2 and 3a11) at lower levels, which mitigates toxicity and buys them time? This would be supported by F2A but not by F3B, at least not for the most important Cyp2e1. A moderate difference is shown for Cyp1a2 expression in F3D, but is that enough to explain the different fates? Or are additional post-transcriptional effects on these Cyps at work?

We thank the reviewer for this important question. We fully agree that the differential susceptibility between mid‑zone and pericentral (PC) hepatocytes is likely rooted in the zonal gradient of cytochrome P450 expression. Our spatial transcriptomics data (Revised Figure 2A) show that mid‑zone hepatocytes express Cyp2e1, Cyp1a2, and Cyp3a11 at levels intermediate between PC and periportal (PP) zones. This intermediate expression may generate sufficient NAPQI to activate stress signaling but not so much as to cause immediate mitochondrial collapse, thus “buying time” for adaptive responses. We also appreciate the reviewer’s observation that Cyp2e1 mRNA levels remain highest in the PC zone even after APAP (Revised Figure 3B). However, mRNA abundance does not necessarily reflect functional protein level. In the PC zone, massive necrosis rapidly compromises cellular integrity; as shown in Revised Figure 3D, Cyp1a2 protein declines sharply around the central vein, and we observed similar degradation for Cyp2e1 (data not shown). Consequently, despite sustained Cyp2e1 transcripts, PC hepatocytes are unable to mount an effective stress response because they are already undergoing cell death. By contrast, mid‑zone hepatocytes retain sufficient metabolic capacity to activate the Atf4‑Chop axis while preserving cellular function.

(2) The evidence presented in support of cell cycle arrest of midlobular hepatocytes is not fully convincing: there is no overt difference in S and G2/M gene scores in F2F; the marker genes used for S phase and G1 to S progression in F2G are unusual. Along these lines, one wonders if spatial transcriptomics confirmed the Ki67 immunostaining results in F1 also for specific zones, not only overall, as shown in F2E?

We thank the reviewer for these important observations. We agree that the current spatial transcriptomics (ST) data alone do not provide sufficiently strong support for this conclusion. The limited sensitivity of ST for detecting rare proliferative events further constrains its utility in this context. At baseline, only ~1% of ST spots are Ki67-positive (Revised Figure S1I), and this fraction becomes even lower during the early phase following APAP injury. As a result, there are insufficient Ki67+ spots to robustly assess zonal distribution using ST, which precludes a reliable spatial validation of proliferation patterns at this resolution. For this reason, our primary evidence for zonal proliferation dynamics relies on Ki67 immunohistochemistry (Revised Figure 1), which provides single-cell resolution and higher sensitivity. These data show a marked reduction in Ki67+ hepatocytes specifically in the midlobular zone at 3-6 hours post-APAP, supporting a transient suppression of proliferation in this region. In addition, we agree that the transcriptional evidence for cell cycle arrest was not strong the S and G2/M scores showed no overt difference, and the gene sets used were suboptimal. We have therefore moved these analyses to the supplement and toned down the claims. We have also clarified this limitation in the manuscript (Revised manuscript, page 18, line 518-524)

(3) The authors conclude in line 364 that halting of proliferation by Btg2 favors survival, which raises the question of whether Btg2 knockout causes death in midlobular hepatocytes in F6K. Data addressing this question, that is, the localization and extent of tissue necrosis and ALT levels after APAP, are missing. The efficiency of the knockout of Btg2 is also not given.

We thank the reviewer for this insightful comment. We have included the missing data. Knockdown of Btg2 using AAV8‑CasRx achieved a moderate (~30%) reduction in Btg2 expression (Revised Figure S5D). Despite this partial efficiency, we observed a significant increase in serum ALT levels (~2‑fold), expansion of necrotic areas (~1.5‑fold), and a marked increase in Ki67+hepatocytes (~1.8‑fold) compared to control mice (Revised Figure 6K–N, Revised manuscript, page 14, line 399-406).

(4) Related to the previous question, the BTG2 immunostaining in F6F is not convincing when compared to F6D. One also wonders if it is necessary to apply APAP to find induction of BTG2 by AAV-Ddit3?

We thank the reviewer for this insightful comment. We have included an inset of the original image to better show BTG2 staining in revised Figure 6F. During our study, we tested BTG2 expression in mice transduced with AAV‑TBG‑EGFP or AAV‑TBG‑BTG2 for three weeks without APAP challenge. We observed that BTG2 in these non‑injured livers was predominantly cytoplasmic (Author response image 2), contrasting with the nuclear localization seen after APAP treatment (Figure 6F). Regarding whether it is necessary to apply APAP to find induction of BTG2 by AAV-Ddit3, we think Ddit3 promotes BTG2 expression (as shown in revised Figure F6F), but APAP is necessary for its nuclear translocation.

Immunohistochemical detection of Btg2 in liver tissue from mice transduced with AAV-TBG-EGFP or AAV-TBG-Btg2 for 3 weeks without APAP treatment.

(5) Related to the previous question, the proposed Atf4-Ddit3 axis is challenged by the lack of midlobular induction of Atf4 in the APAP scRNA-seq data published by another group, presented in S4F and G. Further analysis of AAV-Atf4 samples generated for F5 could address whether it is really Atf4 that acts on Ddit3 in APAP toxicity.

We thank the reviewer for this insightful comment. We agree that Atf4 was not among the top 30 active transcription factors in our initial analysis; however, when we extended the list to the top 50, Atf4 was included. We have therefore updated Revised Figures S4F and G to show the top 50 transcription factors. We also appreciate the reviewer’s suggestion to further investigate whether Atf4 directly acts on Ddit3 in the context of APAP toxicity. While this still shows a less pronounced midlobular enrichment for Atf4 compared with Ddit3, we sought additional evidence for a functional Atf4-Ddit3 link. In primary hepatocytes treated with APAP, we observed nuclear co‑localization of Atf4 and Ddit3 (Author response image 3A) and increased nuclear protein levels of both factors (Author response image 3B), supporting their potential cooperative role. We agree that direct analysis of AAV‑Atf4 samples generated for Figure 5 would provide more definitive evidence; unfortunately, co‑staining for Atf4 and Ddit3 on those tissue sections didn’t work well.

Subcellular localization of Atf4 and Chop in primary hepatocytes following APAP treatment.

(A) Immunofluorescence staining of Atf4 and Chop in primary hepatocytes treated with 10 mM APAP for 6 hours or left untreated (UT). Nuclei were counterstained with DAPI. Scale bar as indicated. (B) Primary hepatocytes were treated with 0, 5, or 10 mM APAP for 6 hours. Cytoplasmic and nuclear fractions were isolated and analyzed by western blot. Lamin B1 and α-Tubulin were used as markers for the nucleus and cytoplasm, respectively

(6) Related to the previous question, the ATF4 immunostaining in F5A doesn't look convincing, with many brown pigments appearing to be outside of the nucleus.

We thank the reviewer for this helpful comment. To better demonstrate ATF4 nuclear localization, we have added enlarged insets of the original representative images in revised Figure 5A. These magnified views more clearly show nuclear ATF4 staining after APAP treatment, addressing the concern about extranuclear signal.

(7) It is not ruled out that AAV expression of Atf4 or Btg2 reduces hepatocyte sensitivity to APAP by affecting the expression of the Cyps needed for activation. In other words, does AAV-Atf4 or AAV-Btg2 change the expression of any of the Cyps relevant to APAP in the 3 weeks before APAP application (F5B)?

We thank the reviewer for raising these important points. We measured basal Cyp2e1 protein levels by western blot in AAV‑EGFP, AAV‑Atf4, and AAV‑Btg2 mice without APAP treatment. Compared to AAV‑EGFP controls, Cyp2e1 expression was modestly reduced in the Atf4 and Btg2 groups, respectively (Revised Figure S5A). Although we assessed protein abundance rather than enzymatic activity directly, Cyp2e1 protein levels under basal conditions generally correlate well with activity. Published studies demonstrate that robust protection against APAP hepatotoxicity typically requires >50% suppression of CYP2E1 activity (PMID: 35145060; PMID: 30151903). The minor reductions we observed are therefore far below the threshold needed to explain the 70–90% decreases in serum ALT conferred by Atf4 or Btg2 overexpression (Revised Figures 5D and 6I). Accordingly, altered CYP2E1 activity is unlikely to represent a significant confound in our model.

(8) It is laudable that the authors tried to extend their findings to humans by using snRNA-seq data from a published study (line 391), but it is unclear why they didn't analyze all 10 patients in that study but instead focused on 2 and stated that this small sample number prevented drawing definitive conclusions and could therefore only be mentioned in the discussion.

We thank the reviewer for this clarification. The analysis mentioned in line 391 originally referred to spatial transcriptomics (ST) data from two ALF patients, not snRNA-seq. For the snRNA-seq dataset, we analyzed all 10 patients, but snRNA-seq lacks spatial resolution and cannot reliably assign zonal identity. We stipulate that snRNA-seq requires viable cells and thus likely excludes necrotic/peri-necrotic areas. Therefore, direct zonal comparison with our ST data was not possible. We have now clarified this in the revised manuscript (Revised manuscript, page 18, line 510-519).

Reviewer #3 (Public Review):

The main concern is that the overexpression of ATF4 and DDIT3 is causing reduced cell death and damage by APAP. This makes it harder to understand if these genes are truly increasing survival or if they are just reducing the injury caused by APAP. It may be better to perform overexpression immediately after, or at the same time as APAP delivery. Alternatively, loss-of-function experiments using AAV-shRNAs against these targets could be useful.

We thank the reviewer for raising this important point. We agree that overexpression prior to APAP administration leaves open the question of whether the observed protection reflects true cytoprotection or simply reduced initiation of injury. To address this, we pursued loss‑of‑function approaches. Due to their very low basal expression, AAV‑shRNA‑mediated knockdown of endogenous Atf4 and Ddit3 proved inefficient. We therefore targeted Btg2, a downstream mediator of Ddit3 that inhibits proliferation. Knockdown of Btg2 resulted in a significant increase in APAP‑induced liver injury, as evidenced by elevated ALT levels and expanded necrotic areas (Revised Figure 6K-N). These results indicate that the ATF4‑DDIT3‑BTG2 axis limits hepatocellular damage, consistent with a protective role. We have clarified this point in the revised manuscript (page 15, line 407-414)

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

(1) Clarify how zones were defined when necrosis disrupted pericentral areas. Provide marker validation across time and whether necrotic spots are excluded or not from zonal analysis.

We thank the reviewer for these insightful suggestions. In this study, under normal conditions (APAP 0h), each liver lobule was divided into three zones based on unbiased gene expression profiles. The PP zone was defined by enrichment of PP signature genes (e.g., Alb, Mup20, Cyp2f2, Pck1, Apoa4). The PC zone was defined by high expression of PC markers (e.g., Gs, Cyp2e1, Oat, Cyp1a2, Apoe). The Mid zone comprised regions with intermediate expression of PC and PP markers and elevated levels of Igfbp2 and Hamp (Revised Figure 2A and S1C). Following APAP-induced injury (3 h, 6 h), the PC zone remained identifiable based on residual enrichment of PC signature genes (e.g., Cyp1a2, Glul) despite necrosis and reduced overall transcription, while PP gene expression remained largely unchanged. The Mid zone was defined as the transcriptional cluster between PC and PP regions exhibiting marked reprogramming, (e.g., Sqstm1, Igfbp1) (Revised Figure 2A and S1C). To validate and quantify our zonation approach, we compared it with classical nine even layers from central vein (CV) to portal vein (PV). Immunostaining and quantification for Cyp2f2 (a PP marker), p62 (the protein product of Sqstm1, a Mid marker during early liver injury), Glutamine Synthetase (GS, the protein product of Glul, a PC marker) further corroborated zone definitions at each time point, showing correspondence of our PC (layers 1–2), Mid (layers 3–6), and PP (layers 7–9) (Revised Figure 2 B-D) (Revised manuscript, page 5, lines 119–131, page 6-7, line 174-182).

(2) Test whether the ISR-Btg2 program applies in other models; even targeted validation via qPCR and IF would be valuable.

We thank the reviewer for this insightful comment. To test whether the proposed mechanism applies to other liver injuries, we employed mouse models of partial hepatectomy (PHx) and carbon tetrachloride (CCl4)-induced acute liver injury. In our CCl4 model (administered intraperitoneally in corn oil, with samples collected 18 h post‑injection), the ISR was activated around injury sites, accompanied by decreased proliferation, as evidenced by increased expression of p‑eIF2α, Atf4, Chop, and Btg2, along with reduced Ki67 expression (Revised Figure S6A–G). In PHx model (examined 24 h after surgery), ISR activation was similarly observed around ischemic injury sites, with increased p‑eIF2α, Atf4, Chop, and Btg2 expression and undetectable Ki67 expression (Revised Figure S7A–G). Together, these additional models suggest that the proposed mechanism may be applicable to other types of liver injury (Revised manuscript, page 15, Line 410-426).

(3) Proliferation quantification in liver sections in Figure 1: how to define the zones and why, at the basal level, there is a high proliferation rate in the mid zone? From Figure 1B-C, all three zones showed decreased hepatocyte proliferation, although the mid zone had a higher baseline. Will the mid-zone stand out by converting to the fold change of Ki-67+ hepatocytes decrease?

We thank the reviewer for these insightful comments. To define the pericentral (PC), mid, and periportal (PP) zones, we adopted the classical nine‑layer model of the hepatic lobule described by Lin et al. (PMID: 29618815). Layers 1–2 were designated as the PC zone, layers 3–6 as the mid zone, and layers 7–9 as the PP zone. For quantitative zonal distribution of protein‑positive nuclei (e.g., Ki67, CHOP, ATF4), we calculated a position index (P.I.) based on distances to the nearest central vein (CV) and portal vein (PV), using the law of cosines: P.I. = (x2 + z2 – y2) / (2z2), where x = distance to CV, y = distance to PV, and z = distance between CV and PV. This quantification method has now been included in the Methods section (Revised manuscript, page 33, line 880-885). Consistent with previous reports (PMID: 33632817; PMID: 33632818), we observed a higher baseline proliferation rate in the mid zone, where approximately 70% of proliferating hepatocytes reside under basal conditions, compared to 10% in zone 1 and 20% in zone 3. However, when analyzing the fold change in Ki-67+ hepatocytes, only Mid and PC region showed significant difference in Ki-67+ hepatocytes, other zones showed no significant differences (as shown in the fold-change results in Author response image 1), indicating that the mid zone does not stand out in the fold change analysis. See Author response image 1.

(4) The authors need to strengthen the causal chain with rescue experiments, e.g., Atf4/Chop overexpression and Btg2 knockdown. Link proliferation suppression to survival/ALT directly.

We thank the reviewer for these constructive comments. Besides existing data from Figure 5 (Atf4 overexpression), we included Btg2 knockdown data in the revised Figure. Knockdown of Btg2 using AAV8‑CasRx achieved a moderate (~30%) reduction in Btg2 expression (Revised Figure S5D). Despite this partial efficiency, we observed a significant increase in serum ALT levels (~2‑fold), expansion of necrotic areas (~1.5‑fold), and a marked increase in Ki67+ hepatocytes (~1.8‑fold) compared to control mice (Revised Figure 6K–N) (Revised manuscript, page 14, lines 399–406).

(5) Transduction efficiency, distribution, and expression levels via the AAV overexpression need to be quantified. Key CYP genes in the APAP metabolic pathway need to be assessed to exclude confounds.

We thank the reviewer for raising these important points. We measured basal Cyp2e1 protein levels by western blot in AAV‑EGFP, AAV‑Atf4, and AAV‑Btg2 mice without APAP treatment. Compared to AAV‑EGFP controls, Cyp2e1 expression was modestly reduced in the Atf4 and Btg2 groups, respectively (Revised Figure S5A). Although we assessed protein abundance rather than enzymatic activity directly, Cyp2e1 protein levels under basal conditions generally correlate well with activity. Published studies demonstrate that robust protection against APAP hepatotoxicity typically requires >50% suppression of CYP2E1 activity (PMID: 35145060; PMID: 30151903). The minor reductions we observed are therefore far below the threshold needed to explain the 70–90% decreases in serum ALT conferred by Atf4 or Btg2 overexpression (Revised Figures 5D and 6I). Accordingly, altered CYP2E1 activity is unlikely to represent a significant confound in our model.

We quantified transduction efficiency by immunohistochemical detection of the respective transgene proteins and determined the percentage of positive hepatocytes. At a dose of 1.2 × 1011 viral genomes per animal, average transduction rates were 32% (EGFP), 18% (Atf4), and 23% (Btg2) (Revised Figure S5B). Individual animal transduction efficiency showed a negative correlation with serum ALT levels (e.g. Pearson r = –0.7681, p = 0.0260 for Atf4; Revised Figure S5C), demonstrating that greater transgene expression associates with stronger protection. Although these average transduction rates appear modest relative to the 70–90% reduction in ALT, this apparent disproportion is consistent with the known tendency of AAV‑TBG vectors to transduce hepatocytes preferentially in the pericentral region—the same zone where APAP‑induced necrosis initiates. Pericentral enrichment of transgene expression could thus provide disproportionate protection by targeting the most vulnerable cells. These data are now included in Revised Figure S5A–C and detailed in the Results (page 14, lines 383–399).

(6) The authors claim that the requirement of the Atf4/Chop at the early stage of APAP injury protects hepatocytes from proliferation for survival. What is the consequence if we remove the protective mechanism?

We thank the reviewer for this insightful question. In our model, early induction of Atf4 and Chop functions as a cell survival checkpoint. Removal of this protective mechanism is predicted to result in two deleterious outcomes: (1) Acute exacerbation of necrosis due to the inability of hepatocytes to manage stress-induced bioenergetic demands, and (2) Impaired long-term regeneration due to depletion of the surviving cell pool. We directly tested the acute prediction (< 24 h) in Author response image 4. We deleted Ddit3 specifically in hepatocytes. Initial attempts using AAV-CasRx failed due to negligible baseline Atf4/Chop expression in healthy liver, preventing effective knockdown. We therefore generated hepatocyte-specific Ddit3 knockout mice (Alb∆Ddit3; Author response image 4B). Immunohistochemistry confirmed APAP-induced Chop induction occurs primarily in the centrilobular zone by 6 h (Author response image 4A). Following a two-dose APAP regimen (Author response image 4C), Alb∆Ddit3 mice displayed significantly larger areas of centrilobular necrosis compared to Ddit3fl/fl controls (Author response image 4D; **p < 0.01). Thus, hepatocyte-intrinsic Chop limits acute APAP injury, consistent with its proposed early protective role.

Hepatocyte-specific deletion of Ddit3 exacerbates APAP-induced liver injury.

(A) Immunohistochemical staining of Chop in liver sections at 0,3 and 6 h post-APAP. Red arrows indicate Chop-positive hepatocytes. Scale bar = 50μm. Quantification of zonal distribution of Chop-positive cells in liver sections at 6 h post-APAP is conducted . The statistic is the percentage of Chop-positive hepatocytes in each layer over the total number of Chop-positive hepatocytes. n=3 mice. (B)The construction, genotyping strategy and genotyping results of Alb∆Ddit3 mice. P: positive control; WT: Wild-type; Neg: Blank control(ddH2O). (C) Schematic figure illustrating the experimental strategy for the administration of two doses of APAP to Ddit3fl/fl and Alb∆Ddit3 mice. (D) H&E staining showing liver morphology from Ddit3fl/fl and Alb∆Ddit3 mice at 6 h post-second dose of APAP. Injured area is outlined by black dashed lines. Scale bars = 200 μm. The percentage of injury area is quantified. n = 3- 4 mice/group. Data are represented as means ± SD; *p < 0.05; **p < 0.01; ***p < 0.001; ****p < 0.0001; ns, not significant.

(7) Is there any human relevance to the sensitivity of APAP injury regarding the Atf4/Chop axis?

We thank the reviewer for this insightful comment. During our study, we analyzed a spatial transcriptomics dataset from APAP patients. In one of two analyzed patients, mid-zone hepatocytes exhibited transcriptional signatures remarkably consistent with our murine findings, including: (1) upregulation of Atf4-Chop pathways, and (2) downregulation of cell proliferation genes (Author response image 5). This suggests that this axis may also be involved in the response to APAP injury in humans. However, given the limited sample size, definitive conclusions cannot be drawn at this stage. We have now included this point in the Discussion section (Revised manuscript, page 18, line 510-519).

Spatial transcriptomics (GSE223561) reveals zonal gene expression changes in APAP patients.

Heatmap of ISR, cell death, and cell cycle gene expression across zonal regions in healthy versus APAP‑treated human livers.

(8) Several IHC stainings have a weak signal and need inserts to zoom in for a clear view of the positive signals. Figure 5A, E, G, and Figure 6D, F.

We thank the reviewer for this observation. We agree that the immunostaining signals for several target genes are relatively weak, which reflects their low endogenous expression levels. To address this, we have included higher-magnification insets in the indicated panels (Revised Figure 5A, E, G and Figure 6D, F) to show the positive signals.

Reviewer #2 (Recommendations for the authors):

(1) What is the functional classification of DEG in F2A based on? GO terms?

We thank the reviewer for this constructive question. The functional classification of differentially expressed genes (DEGs) in F2A is based on Gene Ontology (GO) terms. For each DEG, we retrieved its associated GO annotations across the three main categories (biological process, cellular component, molecular function). In cases where a gene was assigned multiple GO terms, we prioritized the most representative or significantly enriched term for functional interpretation. This clarification has been incorporated into the revised figure legend and the according GO number has been included in the figure.

(3) The rationale for focusing on CHOP is not clear because Ddit3 is not shown in the spatial transcriptomics in F2A and is not significant in F2B, contradicting what is stated in line 206.

We thank the reviewer for raising this important point. We apologize that Ddit3 was missing from the original figure. In the revised manuscript, we have included an updated version of Figure 2A, which now shows that Ddit3 is indeed one of the differentially expressed genes (DEGs) in the Mid zone at both 3 and 6 hours post-APAP. We agree with the reviewer that, as shown in Figure S1G (previous Figure 2B), Ddit3 did not reach statistical significance, due to its relatively low expression level in that analysis. Nevertheless, when we examined transcription factor (TF) activity in the Mid zone during early AILI, Ddit3 and Atf3 ranked as the top two most highly expressed TFs among the top ten with the highest activity, whereas Atf4 ranked seventh (Revised Figure 4B and Figure S3B). Given that Ddit3 frequently co-worked with Atf4 and that the Atf4–Ddit3 axis plays a well-established role in cellular stress adaptation, we considered this pathway to be biologically relevant and worthy of further investigation.

(3) The term "redistribution" used in line 197 to describe the expression of Cyp2e1 and other Cyps in the midlobular zone seems inappropriate, considering that they just continue to be expressed there, whereas pericentral hepatocytes are dying in F3B; the same applies to "Gene Expression Shift" in F3H.

We thank the reviewer for this important clarification. We have revised the text (Revised manuscript, page 9, line 234-236) to state that selective loss of Cyp‑expressing pericentral hepatocytes leads to the mid‑zone becoming the primary site of residual Cyp activity. The figure label has been changed from “Gene Expression Shift” to “Peri‑necrotic Cyp retention” and the legend now explicitly notes that this is an apparent zonal shift due to necrosis, not active redistribution.

Reviewer #3 (Recommendations for the authors):

(1) Please do not use abbreviations like AILI. This makes the paper more difficult to read.

We thank the reviewer for pointing this out. We have replaced AILI with the full term “APAP-induced liver injury” to ensure easiness for readers.

(2) It will be important to clarify how pericentral, mid, and periportal were defined. In Figure 1, it appears that some of the pericentral hepatocytes that are Ki67 positive are quite mid-zonal. It would be important to have rigorous definitions for the location determination.

We thank the reviewer for this constructive comment. To define the pericentral (PC), mid, and periportal (PP) zones, we adopted the classical nine‑layer model of the hepatic lobule described by Lin et al. (PMID: 29618815). Layers 1–2 were designated as the PC zone, layers 3–6 as the mid zone, and layers 7–9 as the PP zone. For quantitative zonal distribution of protein‑positive nuclei (e.g., Ki67, CHOP, ATF4), we calculated a position index (P.I.) based on distances to the nearest central vein (CV) and portal vein (PV), using the law of cosines: P.I. = (x 2 + z 2 – y 2) / (2z 2), where x = distance to CV, y = distance to PV, and z = distance between CV and PV. This quantification method has now been included in the Methods section (Revised manuscript, page 33, line 880-885).

We thank the reviewers for their rigorous critique again. We thank eLife for fostering an environment of fairness and transparency that enables authors to communicate openly and present their data honestly.

Changes from the earlier version to the revised version

Abstract

Earlier version: Hepatocytes undergo extensive proliferation to facilitate liver repair after injury, yet early adaptive changes prior to proliferation remain unclear. Here, we report that during early acetaminophen (APAP)-induced liver injury (AILI), hepatocytes exhibit transient proliferation suppression, most pronounced in mid-zone hepatocytes due to zonal APAP metabolism. Using spatial transcriptomics, immunohistochemistry, and functional studies, we identified a unique mid-zone stress-response program. Central to this adaptation is the Atf4-Chop axis, which actively suppresses proliferation via the cell cycle inhibitor Btg2, prioritizing cytoprotection over cell division. This transient arrest is a critical survival strategy: halting energy-intensive proliferation during peak injury allows mid-zone hepatocytes to redirect resources towards protection, enhancing their survival in early AILI. Thus, Atf4-Chop-mediated quiescence preserves a hepatocyte reservoir necessary for subsequent regenerative proliferation and effective repair. Our findings reveal a key adaptive trade-off in mid-zone hepatocytes where transient proliferation arrest promotes early survival to enable repair.

Revised version: Hepatocytes undergo extensive proliferation to facilitate liver repair after injury, yet early adaptive changes prior to proliferation remain unclear. Here, we report that during early acetaminophen (APAP)-induced liver injury, hepatocytes exhibit transient proliferation suppression, most pronounced in mid-zone hepatocytes due to zonal APAP metabolism. Using spatial transcriptomics (ST), immunohistochemistry, and functional studies, we identified a unique mid-zone stress-response program. Central to this adaptation is the Atf4-Chop axis, which actively suppresses proliferation via the cell cycle inhibitor Btg2, prioritizing cytoprotection over cell division. This transient arrest is a critical survival strategy: halting energy-intensive proliferation during peak injury allows mid-zone hepatocytes to redirect resources towards protection, enhancing their survival in early APAP-induced liver injury. Thus, Atf4-Chop-mediated quiescence preserves a hepatocyte reservoir necessary for subsequent regenerative proliferation and effective repair. Our findings reveal a key adaptive trade-off in mid-zone hepatocytes where transient proliferation arrest promotes early survival to enable repair.

Introduction

Earlier version: Acetaminophen (APAP) overdose is the leading cause of drug-induced acute liver failure worldwide, resulting in over 10,000 hospitalizations and approximately 500 deaths annually in the United States.1–4 The hepatotoxicity of APAP is initiated by its metabolic activation through cytochrome P450 enzymes, primarily in pericentral hepatocytes, generating the reactive metabolite N-acetyl-p-benzoquinone imine (NAPQI).5–8 NAPQI depletes glutathione and forms protein adducts, leading to oxidative stress and mitochondrial dysfunction that culminate in centrilobular necrosis.5–8 Although the mechanisms of APAP bioactivation and toxicity are well-characterized, the early adaptive responses of hepatocytes—particularly their proliferative dynamics across lobular zones—remain poorly understood.

Revised version: Acetaminophen (APAP) overdose is the leading cause of drug-induced acute liver failure worldwide, resulting in over 10,000 hospitalizations and approximately 500 deaths annually in the United States (Bernal and Wendon, 2013; Larson et al., 2005; Lee, 2007; Ostapowicz et al., 2002). The hepatotoxicity of APAP is initiated by its metabolic activation through cytochrome P450 enzymes, primarily in pericentral hepatocytes, generating the reactive metabolite N-acetyl-p-benzoquinone imine (NAPQI) (Dahlin et al., 1984; Hinson et al., 2004; Hinson et al., 2010; Tujios and Fontana, 2011). NAPQI depletes glutathione and forms protein adducts, leading to oxidative stress and mitochondrial dysfunction that culminate in centrilobular necrosis (Dahlin et al., 1984; Hinson et al., 2004; Hinson et al., 2010; Tujios and Fontana, 2011). Although the mechanisms of APAP bioactivation and toxicity are well-characterized, the early adaptive responses of hepatocytes—particularly their proliferative dynamics across lobular zones—remain poorly understood.

Earlier version: The liver is organized into functional hexagonal units termed lobules, where hepatocytes display spatial heterogeneity across three distinct zones: periportal (PP, zone 1), mid-zonal (Mid, zone 2), and pericentral (PC, zone 3).9–11 These zones exhibit unique gene expression profiles and metabolic specializations.9 Pericentral hepatocytes (zone 3), expressing high levels of cytochrome P450 enzymes, are the primary site of acetaminophen (APAP)-induced necrosis.10 However, recent spatial profiling, functional, and lineage tracing studies have identified midzonal hepatocytes (zone 2) as the predominant source of new hepatocytes during both homeostasis and repair, demonstrating significant proliferative capacity and metabolic plasticity.12–14 Despite these advances, how mid-zone hepatocytes dynamically adapt during the initial phase of acute stress remains unknown.

Revised version: The liver is organized into functional hexagonal units termed lobules, where hepatocytes display spatial heterogeneity across three distinct zones: periportal (PP, zone 1), mid-zonal (Mid, zone 2), and pericentral (PC, zone 3) (Ben-Moshe et al., 2022a; Wang et al., 2024; Wu et al., 2024). These zones exhibit unique gene expression profiles and metabolic specializations(Halpern et al., 2017). Pericentral hepatocytes (zone 3), expressing high levels of cytochrome P450 enzymes, are the primary site of APAP-induced necrosis (Wang et al., 2024). However, recent spatial profiling, functional, and lineage tracing studies have identified midzonal hepatocytes (zone 2) as the predominant source of new hepatocytes during both homeostasis and repair, demonstrating significant proliferative capacity and metabolic plasticity (Chembazhi et al., 2021; He et al., 2021; Wei et al., 2021). Despite these advances, how mid-zone hepatocytes dynamically adapt during the initial phase of acute stress remains unknown.

Earlier version: A key regulator of cellular stress adaptation is the integrated stress response (ISR), which phosphorylates eukaryotic initiation factor 2α (eIF2α) to suppress global translation while selectively upregulating stress-responsive genes, including activating transcription factor 4 (Atf4) and its downstream target pro-apoptotic factor C/EBP homologous protein (Chop, encoded by DNA damage-inducible transcript 3 [Ddit3]).15,16 Although Chop is traditionally associated with apoptosis, its transcriptional targets exhibit significant overlap with those of Atf4, encompassing genes that paradoxically support cell survival and growth.17–19 These findings underscore the context-dependent duality of the Atf4-Chop axis across different stress conditions, highlighting the need to elucidate its precise regulatory roles in specific physiological settings.

Revised version: A key regulator of cellular stress adaptation is the integrated stress response (ISR), which phosphorylates eukaryotic initiation factor 2α (eIF2α) to suppress global translation while selectively upregulating stress-responsive genes, including activating transcription factor 4 (Atf4) and its downstream target pro-apoptotic factor C/EBP homologous protein (Chop, encoded by DNA damage-inducible transcript 3 [Ddit3]) (Costa-Mattioli and Walter, 2020; Vasudevan et al., 2020). Although Chop is traditionally associated with apoptosis, its transcriptional targets exhibit significant overlap with those of Atf4, encompassing genes that paradoxically support cell survival and growth (Han et al., 2013; Marciniak et al., 2004; Uzi et al., 2013). These findings underscore the context-dependent duality of the Atf4-Chop axis across different stress conditions, highlighting the need to elucidate its precise regulatory roles in specific physiological settings.

Earlier version: Through spatial transcriptomics (ST) and functional analyses, we demonstrate that zonal heterogeneity in APAP metabolism drives a stress-adaptive response in mid-zone hepatocytes during early AILI. This response is mediated by the Atf4-Chop-Btg2 axis, which orchestrates a transient proliferation arrest to prioritize cellular survival over regenerative capacity. Our findings elucidate the spatial and molecular mechanisms governing hepatocyte adaptation to acute injury, uncovering a critical survival-regeneration trade-off in early phase of acute liver injury.

Revised version: Through spatial transcriptomics (ST) and functional analyses, we demonstrate that zonal heterogeneity in APAP metabolism drives a stress-adaptive response in mid-zone hepatocytes during early APAP-induced liver injury. This response is mediated by the Atf4-Chop-Btg2 axis, which orchestrates a transient proliferation arrest to prioritize cellular survival over regenerative capacity. Our findings elucidate the spatial and molecular mechanisms governing hepatocyte adaptation to acute injury, uncovering a critical survival-regeneration trade-off in early phase of acute liver injury.

Results

Mid-zone hepatocytes exhibit the most pronounced proliferative decline during early acute liver injury

Earlier version: To assess hepatocyte proliferation in early AILI, wild-type C57BL/6J mice were intraperitoneally injected with 300 mg/kg APAP, and liver sections were analyzed at multiple time points (0, 3, 6, 12, and 24 hours post-APAP). Immunohistochemical staining for Ki67, a proliferation marker, revealed that mid-zone hepatocytes exhibited a roughly 3-fold higher basal proliferative activity than PP and PC hepatocytes under homeostatic conditions (0 hour) (Figure 1A, B). During the early injury phase (3–12 hours post-APAP), the number of Ki67+ hepatocytes in the mid-zone sharply declined from 70 to 20 per section, reaching proliferation levels similar to those in PP and PC zones (Figure 1A, B). By 24 hours post-APAP, mid-zone proliferation gradually recovered (Figure 1A, B), indicating a transient suppression of proliferation in this region during early AILI. Further analysis of proliferation dynamics in each zone confirmed that all zones exhibited reduced proliferation during the injury initiation phase, with the mid-zone showing the most decline (Figure 1C). Comparative value change analysis (relative to 0 hour) further supported these findings, highlighting the mid-zone as the most affected region in terms of proliferative suppression during early AILI (Figure 1D).

Revised version: To assess hepatocyte proliferation in early APAP-induced liver injury, wild-type C57BL/6J mice were intraperitoneally injected with 300 mg/kg APAP, and liver sections were analyzed at multiple time points (0, 3, 6, 12, and 24 hours post-APAP). Immunohistochemical staining for Ki67, a proliferation marker, revealed that mid-zone hepatocytes exhibited a roughly 3-fold higher basal proliferative activity than PP and PC hepatocytes under homeostatic conditions (0 hour) (Figure 1A, B), which in line with previous study (Wei et al., 2021). During the early injury phase (3–12 hours post-APAP), the number of Ki67+ hepatocytes in the mid-zone sharply declined from 70 to 20 per section, reaching proliferation levels similar to those in PP and PC zones (Figure 1A, B). By 24 hours post-APAP, mid-zone proliferation gradually recovered (Figure 1A, B), indicating a transient suppression of proliferation in this region during early APAP-induced liver injury. Further analysis of proliferation dynamics in each zone confirmed that all zones exhibited reduced proliferation during the injury initiation phase, with the mid-zone showing the most decline (Figure 1C). Comparative value change analysis (relative to 0 hour) further supported these findings, highlighting the mid-zone as the most affected region in terms of proliferative suppression during early APAP-induced liver injury (Figure 1D).

Mid-zone hepatocytes exhibit the most pronounced proliferative decline in early acute liver injury.

Earlier version: Wildtype C57BL/6J mice were intraperitoneally injected with 300 mg/kg APAP for 0, 3, 6, 12 and 24 h. (A) Immunohistochemical staining of Ki67 to detect proliferating hepatocytes in liver sections at 0, 3, 6, 12 and 24 h post-APAP. Representative images of pericentral (PC), middle (Mid) and periportal (PP) are shown. Red arrows indicate Ki67-positive hepatocytes. (B) Quantitative analysis of Ki67+ hepatocytes across liver zones (PC, Mid, PP) at 0, 3, 6, 12, and 24 h post-APAP. (C) Quantitative analysis of Ki67+ hepatocytes across 0, 3, 6, 12, and 24 h post-APAP in each liver zone (PC, Mid, PP). (D) Changes of Ki67-positive cells across liver zones (PC, Mid, PP) at 0, 3, 6, 12 and 24 h post-APAP. Data are presented as log2-value change (number of Ki67-positive cells at x h post-APAP – the number of Ki67-positive cells at 0 h post-APAP). (a) denotes significance between PC and Mid regions, (b) denotes significance between PC and PP regions, and (c) denotes significance between Mid and PP regions. Data points represent mean hepatocyte counts from six high-power fields per zonal layer (n=3 mice/group) in B and C. Data are represented as means ± SD; *p < 0.05; **p < 0.01; ***p < 0.001; ****p < 0.0001; ns, not significant.

Revised version: Wildtype C57BL/6J mice were intraperitoneally injected with 300 mg/kg APAP for 0, 3, 6, 12 and 24 h. (A) Immunohistochemical staining of Ki67 to detect proliferating hepatocytes in liver sections at 0, 3, 6, 12 and 24 h post-APAP. Representative images of pericentral (PC), middle (Mid) and periportal (PP) are shown. Red arrows indicate Ki67+ hepatocytes. (B) Quantitative analysis of Ki67+ hepatocytes across liver zones (PC, Mid, PP) at 0, 3, 6, 12, and 24 h post-APAP. (C) Quantitative analysis of Ki67+ hepatocytes across 0, 3, 6, 12, and 24 h post-APAP in each liver zone (PC, Mid, PP). (D) Changes of Ki67+ hepatocytes across liver zones (PC, Mid, PP) at 0, 3, 6, 12 and 24 h post-APAP. Value change=number of Ki67+ hepatocytes at x h post-APAP - the number of Ki67+ hepatocytes at 0 h post-APAP. (a) denotes significance between PC and Mid regions, (b) denotes significance between PC and PP regions, and (c) denotes significance between Mid and PP regions. Data points represent mean hepatocyte counts from six high-power fields per zonal layer (n=3 mice/group) in B and C. Data are represented as means ± SD; One-way ANOVA (B,C,D).*p < 0.05; **p < 0.01; ***p < 0.001; ****p < 0.0001; ns, not significant.

Earlier version: Mid-zone hepatocytes show distinct gene expression profiles with reduced proliferative activity during early acute injury

Revised version: Mid-zone hepatocytes show distinct gene expression profiles

Earlier version: To investigate the mechanism underlying proliferation pause during early AILI, liver sections from wild-type mice treated with APAP for various time points were subjected to ST using the Visium (10 × Genomics) platform (Figure S1A). Due to variability in mice’s responses to APAP, multiple mice were treated at each time point, and the samples representing the average level at each time point were selected based on Hematoxylin and eosin (H&E) staining and serum alanine aminotransferase (ALT) levels for ST analysis. H&E staining of liver sections used for ST revealed a gradual increase in the necrotic area from 0 to 24 hours (Figure S1B). Spatial map analysis consistently and unbiasedly divided the liver into three zones (PP, Mid, and PC) across all time points (Figures S1C). Analysis of ST data at 0, 3, 6, and 24 hours post-APAP via UMAP plots consistently revealed distinct zonal clusters and dynamic spatiotemporal patterns during the early phase of acute injury (Figures S1D and S1E). Spatiotemporally resolved heatmaps of representative markers, such as Glul (a PC marker) and Cyp2f2 (a PP marker), further reinforced the stability of zonation profiles during early acute liver injury (Figure S1F). These findings suggest that while gene expression profiles change during acute injury, the fundamental spatial architecture of liver zones remains stable. Additionally, immunofluorescence staining of Glutamine Synthetase (GS, the protein product of Glul) and RNA-FISH for Cyp2f2 in liver sections collected at 6 hours post-APAP confirmed these ST analyses (Figure S1G).

Revised version: To investigate the mechanism underlying proliferation pause during early APAP-induced liver injury, liver sections from wild-type mice treated with APAP for various time points were subjected to ST using the Visium (10 × Genomics) platform (Figure S1A). Due to variability in mice’s responses to APAP, multiple mice were treated at each time point and one sample representing the average level at each time point was selected based on Hematoxylin and eosin (H&E) staining and serum alanine aminotransferase (ALT) levels for ST analysis. H&E staining of liver sections used for ST revealed a gradual increase in the necrotic area from 0 to 24 hours (Figure S1B). Under normal conditions (APAP 0h), each liver lobule was divided into three zones based on unbiased gene expression profiles. The PP zone was defined by enrichment of PP signature genes (e.g., Alb, Mup20, Cyp2f2, Pck1, Apoa4). The_PC zone was defined by high expression of PC markers (e.g., Glul, Cyp2e1, Oat, Cyp1a2, Apoe). The Mid zone comprised regions with intermediate expression of PC and PP markers and elevated levels of Igfbp2 and Hamp (Figure 2A and S1C). Following APAP-induced injury (3h, 6h), the PC zone remained identifiable based on residual enrichment of PC signature genes (e.g., Cyp1a2, Glul) despite necrosis and reduced overall transcription, while PP gene expression remained largely unchanged. The Mid zone was defined as the transcriptional cluster between PC and PP regions exhibiting marked reprogramming, (e.g., Sqstm1, Igfbp1) (Figure 2A and S1C). Analysis of ST data at 0, 3, 6, and 24 hours post-APAP via UMAP plots consistently revealed distinct zonal clusters and dynamic spatiotemporal patterns during the early phase of acute injury (Figures S1D and S1E). Spatiotemporally resolved heatmaps of representative markers, such as Glul (a PC marker) and Cyp2f2 (a PP marker), further reinforced the stability of zonation profiles during early acute liver injury (Figure S1F).

Removed in the revised version: We further analyzed differentially expressed genes (DEGs) in hepatocyte zones (PP, Mid, and PC) at 0, 3, 6, and 24 hours post-APAP (Figure 2A). Throughout the 0 to 24-hour period post-APAP, gene expression in the PP remained relatively stable. At 0, 3, and 6 hours post-APAP, genes associated with metabolism (such as Hsd17b13, Fabp1, Apoc2, Thrsp for fat metabolism; Rida, Hal, Gls2, Bhmt for fatty acid metabolism; Aldob, Pck1 for glucose metabolism) and transport (such as Alb, Hpx, Selenop, Trf, Lcn2) were prominently expressed, with a notable upregulation of inflammatory response-related genes (such as Orm1, Serpina1c, Lrg1) observed at 24 hours post-APAP (Figure 2A). In contrast, Mid-zone exhibited significant changes over time. Initially, it showed high expression of genes related to cell proliferation (such as Igfbp2), metabolism (such as Scd1, Cdo1, Ces1c, Ttr, Apoa1), and homeostasis (such as Hamp, Cps1, Rgn) at 0 hours, transitioning to stress response/cell survival (such as Hmox1, Hspa8, Atf3, Osgin1) and protein degradation/autophagy (such as Ubb, Ubc, Sqstm1) at 3 and 6 hours, and eventually shifting to cell proliferation and wound healing (such as Fgg, Fgl1, Fgb) by 24 hours post-APAP (Figure 2A). The PC showed a decrease in xenobiotic metabolism (such as Cyp3a11, Cyp2e1, Cyp1a2, Cyp2c37) and fat metabolism (such as Apoe, Akr1c6, Hpd, Oat) during early AILI (at 0, 3, 6 hours post-APAP), with a marked increase in the expression of genes related to cytoskeleton (such as Actb, Krt8, Krt18) and redox regulation (such as Prdx1, Srxn1) observed at 24 hours post-APAP (Figure 2A). These differential expression patterns across various zones and time points post-APAP highlight significant changes in Mid-zone. Notably, our analysis suggests that the Mid-zone undergoes a stress adaption during the initial phases of acute liver injury.

Earlier version: Mid-zone hepatocytes show distinct gene expression profiles with reduced proliferative activity during early acute injury.

Revised version: Mid-zone hepatocytes show distinct gene expression profiles.

Added in the revised version: (A) Heatmap of the top 10 differentially expressed genes in periportal (PP), midzonal (Mid) and pericentral (PC) hepatocyte zones at 0, 3, 6 and 24 h post-APAP. Functional categories at the bottom were assigned using Gene Ontology (GO) terms (biological process, cellular component, molecular function); for genes annotated with multiple GO terms, the most representative or significantly enriched term was selected. (B) Immunofluorescence staining for glutamine synthetase (GS, a pericentral marker) with DAPI (blue) in liver sections at 0, 3, 6, and 24 h after APAP administration. Quantification of the zonal distribution of GS-positive cells at the indicated time points. (C) Immunohistochemical detection of p62 (Sqstm1 protein, a midzone marker during early injury) in liver sections at 0, 3, 6, and 24 h after APAP administration. Red arrows indicate representative p62-positive hepatocytes. Quantification of the zonal distribution of p62-positive cells at the indicated time points. (D) Immunohistochemical staining of Cyp2f2 (a periportal marker) in liver sections at 0, 3, 6, and 24 h after APAP administration. Quantification of the zonal distribution of Cyp2f2-positive cells at the indicated time points. Data are presented as mean ± SD. One-way ANOVA (B-D). *p < 0.05, **p < 0.01, ***p < 0.001, ****p < 0.0001; ns, not significant.

Added in the revised version: We further analyzed differentially expressed genes (DEGs) in hepatocyte zones (PP, Mid, and PC) at 0, 3, 6, and 24 hours post-APAP (Figure 2A). Throughout the 0 to 24-hour period post-APAP, gene expression in the PP remained relatively stable. At 0, 3, and 6 hours post-APAP, genes associated with metabolism (such as Hsd17b13, Fabp1, Apoc2, Thrsp for fat metabolism; Rida, Hal, Gls2, Bhmt for fatty acid metabolism; Aldob, Pck1 for glucose metabolism) and transport (such as Alb, Hpx, Selenop, Trf, Lcn2) were prominently expressed, with a notable upregulation of inflammatory response-related genes (such as Orm1, Serpina1c, Lrg1) observed at 24 hours post-APAP (Figure 2A). In contrast, Mid-zone exhibited significant changes over time. Initially, it showed high expression of genes related to cell proliferation (such as Igfbp2), metabolism (such as Scd1, Cdo1, Ces1c, Ttr, Apoa1), and homeostasis (such as Hamp, Cps1, Rgn) at 0 hours, transitioning to stress response/cell survival (such as Hmox1, Hspa8, Atf3, Osgin1) and protein degradation/autophagy (such as Ubb, Ubc, Sqstm1) at 3 and 6 hours, and eventually shifting to cell proliferation and wound healing (such as Fgg, Fgl1, Fgb) by 24 hours post-APAP (Figure 2A). The PC showed a decrease in xenobiotic metabolism (such as Cyp3a11, Cyp2e1, Cyp1a2, Cyp2c37) and fat metabolism (such as Apoe, Akr1c6, Hpd, Oat) during early APAP-induced liver injury (at 0, 3, 6 hours post-APAP), with a marked increase in the expression of genes related to cytoskeleton (such as Actb, Krt8, Krt18) and redox regulation (such as Prdx1, Srxn1) observed at 24 hours post-APAP (Figure 2A). These differential expression patterns across various zones and time points post-APAP highlight significant changes in Mid-zone. Notably, our analysis suggests that the Mid-zone undergoes a stress adaption during the initial phases of acute liver injury.

Removed in the revised version: (A) Heatmap showing the gene expression of the top ten differentially expressed genes in hepatocyte zones (zones PP, Mid, and PC) at 0, 3, 6, 24 h post-APAP. The major functions of hepatocytes in each zone are labeled on the left side for 0, 6, and 24 h post-APAP, respectively. (B) Volcano plot illustrating the DEGs in Mid-zone at 6 h compared to 0 h post-APAP. Gray dots denote genes that are not statistically significant. Red and blue dots represent genes that are upregulated and downregulated, respectively, in the sample tissue, at least a 0.5-fold difference from the matched control, with a false discovery rate (FDR) threshold of 0.05. (C) Scatter plot shows the dynamic changes of Sqstm1 mean expression level in hepatocyte zones (zones PP, Mid, and PC) at 0 and 6 h post-APAP. (D) Immunohistochemical staining of p62 (protein product of Sqstm1) to detect the dynamic changes of p62 across hepatocyte zones at 0 and 6 h post-APAP. Red arrows indicate p62-positive hepatocytes. (E) Percentage of Ki67-positive spots analyzed at 0, 3, 6 and 24 h post-APAP. (F) Average module score for S-phase (up) and G2/M-phase (down) genes across hepatocyte zones (zones PP, Mid, and PC) at 0, 3, 6, 24 h post-APAP. (G) Expression levels of the S-phase gene Nasp (up) and the G2/M phase gene Cks1b (down) in each hepatocyte zone (zones PP, Mid, and PC) at 0, 3, 6, 24 h post-APAP.

Earlier version: Further analysis focuses on DEGs in Mid-zone at 6 hours post-APAP compared to 0 hour. The volcano plot revealed a notable increase in the number of genes significantly upregulated in the mid zones at 6 hours, predominantly associated with the stress response and adaptation (such as Atf3, Hmox1, Osgin1, Ddit3, Gadd45a, Atf4, Jun, Hspa1) as well as cell proliferation and apoptosis (such as Btg2, Egr1, Igfbp1, Sqstm1, Dnajb1) (Figure 2B). Among the significantly upregulated genes, Sqstm1 (Sequestosome 1) was selected for further validation. Sqstm1 encodes the p62 protein, a key regulator of autophagy and oxidative stress response.20 While Sqstm1 expression was almost evenly distributed across all zones at 0 hours, it was markedly elevated in the mid-zone at 6 hours post-APAP (Figure 2C). Immunohistochemical staining further confirmed the distinct localization of p62 in the mid-zone at this time point (Figure 2D). Together, our findings indicate that Mid-zone hepatocytes undergo dynamic changes during the early phases of acute liver injury, primarily characterized by stress response and adaptation.

Revised version: Further analysis focuses on DEGs in Mid-zone at 6 hours post-APAP compared to 0 hour. The volcano plot revealed a notable increase in the number of genes significantly upregulated in the mid zones at 6 hours, predominantly associated with the stress response and adaptation (such as Atf3, Hmox1, Osgin1, Ddit3, Gadd45a, Atf4, Jun, Hspa1) as well as cell proliferation and apoptosis (such as Btg2, Egr1, Igfbp1, Sqstm1, Dnajb1) (Figure S1G). Among the significantly upregulated genes, Sqstm1 (Sequestosome 1) was selected for further validation. Sqstm1 encodes the p62 protein, a key regulator of autophagy and oxidative stress response (Huang et al., 2023). While Sqstm1 expression was almost evenly distributed across all zones at 0 hours, it was markedly elevated in the mid-zone at 3 and 6 hours post-APAP (Figure S1H). To validate and quantify our zonation approach during early injury, we compared the assigned zones with a classical division of the lobule into nine even layers spanning from the central vein (CV) to the portal vein (PV). Immunostaining and quantification for glutamine synthetase (GS, the protein product of Glul, a PC marker) and immunohistochemical staining for p62 (a midzone marker during early injury) and Cyp2f2 (a PP marker) further corroborated the zone definitions at each time point, demonstrating that our PC, Mid, and PP zones corresponded to layers 1–2, 3–6, and 7–9, respectively (Figure 2B–D). These findings suggest that while gene expression profiles change during acute injury, the fundamental spatial architecture of liver zones remains stable.

Earlier version: Further ST analysis revealed that Ki67-positive spots also show a transient reduction during the early injury (at 3 and 6 hours post-APAP) (Figure 2E), which aligns with the findings from Ki67 immunohistochemical staining. Cell cycle analysis is crucial for understanding cell proliferation, as it reviews the stages of the cell cycle that cells are in and how many cells are actively dividing.(23) In the cell cycle, the S phase is responsible for DNA synthesis and replication, while the G2M phase prepares cells for and executes mitosis.(23) Analysis of cell-cycle phase distribution, based on genes identified in previous studies,(24, 25) revealed that Mid-zone hepatocytes exhibited a transient increase in S-phase genes at 3 hours, followed by a decline at later time points (6, and 24 hours). In contrast, Mid-zone hepatocytes exhibited reduced G2/M-phase scores at 3 and 6 hours, returning to baseline levels by 24 hours (Figure 2F). This suggests that Mid-zone hepatocytes may initiate DNA synthesis, but they are unable to complete the cycle due to the environmental toxicity or cellular damage. Nuclear Autoantigenic Sperm Protein (Nasp) binds to histones and facilitate chromatin assembly during the S phase.21 Casein Kinase 1 Beta (Cks1b) participates in the control of G1 phase and the transition to the S phase.22 High Nasp levels indicate efforts toward DNA repair and chromatin stability, while low Cks1b levels point to a potential cell cycle arrest or a reduction in cell proliferation to manage the damage (Figure 2G). Overall, ST analysis revealed that Mid-zone hepatocytes exhibit distinct gene expression profiles and cell cycle arrest during early acute injury.

Revised version: Further ST analysis revealed that Ki67-positive spots also show a transient reduction during the early injury (at 3 and 6 hours post-APAP) (Figure S1I), which aligns with the findings from Ki67 immunohistochemical staining. Cell cycle analysis is crucial for understanding cell proliferation, as it reveals the stages of the cell cycle that cells are in and how many cells are actively dividing (Woo et al., 2009). In the cell cycle, the S phase is responsible for DNA synthesis and replication, while the G2M phase prepares cells for and executes mitosis (Woo et al., 2009). Analysis of cell-cycle phase distribution, based on genes identified in previous studies (Bunch et al., 2023; Ohoka et al., 2005), revealed that Mid-zone hepatocytes exhibited a transient increase in S-phase genes at 3 hours, followed by a decline at later time points (6, and 24 hours). In contrast, Mid-zone hepatocytes exhibited reduced G2/M-phase scores at 3 and 6 hours, returning to baseline levels by 24 hours (Figure S1J). This suggests that Mid-zone hepatocytes may initiate DNA synthesis, but they are unable to complete the cycle due to the environmental toxicity or cellular damage. Nuclear Autoantigenic Sperm Protein (Nasp) binds to histones and facilitate chromatin assembly during the S phase (Richardson et al., 2006). Casein Kinase 1 Beta (Cks1b) participates in the control of G1 phase and the transition to the S phase (Peters et al., 2000). High Nasp levels indicate efforts toward DNA repair and chromatin stability, while low Cks1b levels point to a potential cell cycle arrest or a reduction in cell proliferation to manage the damage (Figure S1K). Overall, ST analysis revealed that Mid-zone hepatocytes exhibit distinct gene expression profiles during early acute injury.

Zonal metabolism of APAP leads to stress response in Mid-zone hepatocytes

Earlier version: Next, we aimed to elucidate the cause for the stress response observed in the Mid-zone during early APAP. APAP is metabolized by cytochrome P450 enzymes (Cyp family) into reactive metabolites, particularly N-acetyl-p-benzoquinone imine (NAPQI), which can induce oxidative stress and damage cellular components if not promptly detoxified by glutathione (GSH) (Figure 3A). We analyzed the gene expression levels of Cyp family enzymes in hepatocyte zones (PP, Mid, and PC) at 0, 3, 6, and 24 hours post-APAP. Our analysis revealed a gradient decrease in the expression levels of most Cyp enzymes, showing a trend from high expression in PC to lower expression in PP (Figure 3B). Scatter plots of representative Cyp enzymes, such as Cyp2e1 and Cyp1a2, clearly showed a temporal and spatial expression pattern consistent with this gradient decrease from CV to PV (Figure S2A). Immunohistochemical staining of Cyp1a2 confirmed its gradient decrease in protein expression from the CV to the PV under homeostatic conditions (0 hour). However, following APAP-induced toxicity, Cyp1a2 expression declined around the CV— where necrosis was prominent—during the injury initiation phase (3–24 hours post-APAP). Consequently, the Mid-zone adjacent to necrotic regions emerged as the primary area retaining Cyp1a2 expression (Figure 3C, D). These findings suggest that the spatial redistribution of Cyp expression establishes the Mid-zone as the primary site of APAP metabolism during early AILI, thereby inducing subsequent transcriptional reprogramming (Figure 3E).

Revised version: Next, we aimed to elucidate the cause for the stress response observed in the Mid-zone during early APAP-induced liver injury. APAP is metabolized by cytochrome P450 enzymes (Cyp family) into reactive metabolites, particularly N-acetyl-p-benzoquinone imine (NAPQI), which can induce oxidative stress and damage cellular components if not promptly detoxified by glutathione (GSH) (Figure 3A). We analyzed the gene expression levels of Cyp family enzymes in hepatocyte zones (PP, Mid, and PC) at 0, 3, 6, and 24 hours post-APAP. Our analysis revealed a gradient decrease in the expression levels of most Cyp enzymes, showing a trend from high expression in PC to lower expression in PP (Figure 3B). Scatter plots of representative Cyp enzymes, such as Cyp2e1 and Cyp1a2, clearly showed a temporal and spatial expression pattern consistent with this gradient decrease from CV to PV (Figure S2A). Immunohistochemical staining of Cyp1a2 confirmed its gradient decrease in protein expression from the CV to the PV under homeostatic conditions (0 hour). However, following APAP-induced toxicity, Cyp1a2 expression declined around the CV—where necrosis was prominent—during the injury initiation phase (3–24 hours post-APAP). Consequently, the Mid-zone adjacent to necrotic regions emerged as the primary area retaining Cyp1a2 expression (Figure 3C, D). These results suggest that selective loss of Cyp- expressing pericentral hepatocytes leads to the mid- zone becoming the primary site of residual Cyp activity (Figure 3E).

Zonal metabolism of APAP leads to stress response in Mid-zone hepatocytes.

Earlier version: (A) Schematic figure illustrating the metabolic process of liver injury caused by APAP. (B) Dot plot displays the scaled gene expression levels of the Cyp family in hepatocyte zones (zones PP, Mid, and PC) at 0, 3, 6, 24 h post-APAP. (C) Immunohistochemical staining of Cyp1a2 to detect its dynamic changes in hepatocyte zones (from CV to PV) at 0, 3, 6, 24 h post-APAP. Injured area is outline by black dashed lines. (D) Quantification of zonal distribution of Cyp1a2-positive cells in liver sections at 0, 3, 6, and 24 h post-APAP, respectively. (E) Schematic figure illustrating expression changes of Cyp1a2 before and after APAP. (F) Immunofluorescence staining of Chop protein (green) to observe the dynamic changes in its expression across hepatocyte zones (from CV to PV) at 6 h post-APAP with doses of 100, 300, or 500 mg/kg. Cyp1a2 protein (red) staining highlights the area around the CV. Cell nuclei are stained with DAPI (blue). (G) Quantification of zonal distribution of Chop-positive cells in liver sections from mice treated with varying doses of APAP– 100, 300, or 500 mg/kg. (H) Schematic figure illustrating zonal metabolism of APAP results in the gene expression shift in Mid hepatocytes following necrosis. Data are represented as means ± SD; *p < 0.05; **p < 0.01; ***p < 0.001; ****p < 0.0001; ns, not significant.

Revised version: (A) Schematic figure illustrating the metabolic process of liver injury caused by APAP. (B) Dot plot displays the scaled gene expression levels of the Cyp family in hepatocyte zones (zones PP, Mid, and PC) at 0, 3, 6, 24 h post-APAP. (C) Immunohistochemical staining of Cyp1a2 to detect its dynamic changes in hepatocyte zones (from CV to PV) at 0, 3, 6, 24 h post-APAP. Injured area is outline by black dashed lines. (D) Quantification of zonal distribution of Cyp1a2-positive cells in liver sections at 0, 3, 6, and 24 h post-APAP, respectively. (E) Schematic figure illustrating expression changes of Cyp1a2 before and after APAP. (F) Immunofluorescence staining of Chop protein (green) to observe the dynamic changes in its expression across hepatocyte zones (from CV to PV) at 6 h post-APAP with doses of 100, 300, or 500 mg/kg. Cyp1a2 protein (red) staining highlights the area around the CV. Cell nuclei are stained with DAPI (blue). (G) Quantification of zonal distribution of Chop-positive cells in liver sections from mice treated with varying doses of APAP– 100, 300, or 500 mg/kg. (H) Schematic explaining why mid-lobular hepatocytes display a distinct gene expression program early after APAP injury. Cyp enzymes are expressed in a decreasing gradient from the central vein (CV) to the portal vein (PV). APAP is therefore metabolized primarily by pericentral Cyp-expressing hepatocytes; their necrosis shifts the main site of residual Cyp activity to the mid-lobular zone. High local APAP metabolism then drives a specific stress-response program, accounting for the unique early transcriptional signature of the mid-zone. Data are represented as means ± SD; One-way ANOVA (D,G). *p < 0.05; **p < 0.01; ***p < 0.001; ****p < 0.0001; ns, not significant.

Earlier version: To investigate this hypothesis, we administered mice with varying doses of APAP (100, 300, and 500 mg/kg) to induce different extents of liver injury. H&E staining revealed a dose-dependent increase in the size of the injured areas, with higher doses resulting in larger areas of injury (Figure S2B and S2C). Subsequently, we performed immunofluorescence staining for Chop, a protein differentially highly expressed in the Mid-zone at 6 hours post-APAP compared to 0 hour (as identified in Figure 2B), to assess the Mid-zone response. Our findings indicate that Chop expression is absent at low doses of APAP when necrosis does not occur. However, as the dose of APAP increases, the injured area expands, and Chop expression shifts toward the PV, consistently surrounding the injured area (Figure 3F and G). Together, these findings demonstrate that CV necrosis establishes the adjacent peri-necrotic region as the dominant site of APAP metabolism, contingent upon both Cyp family expression levels and APAP dosage. This spatially restricted metabolic adaptation drives subsequent region-specific transcriptional changes in the surrounding parenchyma (Figure 3H).

Revised version: To investigate this hypothesis, we administered mice with varying doses of APAP (100, 300, and 500 mg/kg) to induce different extents of liver injury. H&E staining revealed a dose-dependent increase in the size of the injured areas, with higher doses resulting in larger areas of injury (Figure S2B and S2C). We then performed immunofluorescence staining for Chop, the protein product of Ddit3, which was differentially upregulated in the Mid-zone at 6 hours post-APAP compared to 0 hours (as shown in Figure 2B), to assess the Mid-zone response. Our findings indicate that Chop expression is absent at low doses of APAP when necrosis does not occur. However, as the dose of APAP increases, the injured area expands, and Chop expression shifts toward the PV, consistently surrounding the injured area (Figure 3F and G). Together, these findings demonstrate that CV necrosis establishes the adjacent peri-necrotic region as the dominant site of APAP metabolism, contingent upon both Cyp family expression levels and APAP dosage. This spatially restricted metabolic adaptation drives subsequent region-specific transcriptional changes in the surrounding parenchyma (Figure 3H).

The Atf4-Chop axis emerges as pivotal in Mid-zone during early acute injury

Earlier version: To explore master regulator of stress adaptation in Mid-zone, we constructed gene regulatory networks from the ST data. We identified the top 10 regulons in each zone at 3 and 6 hours post-APAP (Figure 4A and Figure S3A). A closer examination of the transcription factors (TFs) in the Mid-zone revealed that Atf4, Fos, Nfil3, Egr1, Maff, Atf3, and Ddit3 co-activated at both time points. Atf4, Chop, and Atf3 play crucial roles in managing cellular stress responses.23,24 Fos and Egr1 are immediate early genes that are rapidly activated in response to stimuli, regulating genes related to growth and stress responses.25 Maff and Nfil3 contribute to the management of oxidative stress and metabolic processes.26,27 The elevated activity of these TFs in the Mid-zone indicates that this region is undergoing a complex response to stress, serving as a critical area for cellular adaptation, survival, and potential necrosis during early liver injury.

Revised version: To explore master regulator of stress adaptation in Mid-zone, we constructed gene regulatory networks from the ST data. We identified the top 10 regulons in each zone at 3 and 6 hours post-APAP (Figure 4A and Figure S3A). A closer examination of the transcription factors (TFs) in the Mid-zone revealed that Atf4, Fos, Nfil3, Egr1, Maff, Atf3, and Ddit3 co-activated at both time points. Atf4, Chop, and Atf3 play crucial roles in managing cellular stress responses (Ohoka et al., 2005; Woo et al., 2009). Fos and Egr1 are immediate early genes that are rapidly activated in response to stimuli, regulating genes related to growth and stress responses (Bunch et al., 2023). Maff and Nfil3 contribute to the management of oxidative stress and metabolic processes (Lin et al., 2024; von Scheidt et al., 2021). The elevated activity of these TFs in the Mid-zone indicates that this region is undergoing a complex response to stress, serving as a critical area for cellular adaptation, survival, and potential necrosis during early liver injury.

The Atf4-Ddit3 axis emerges as pivotal in Mid-zone during early acute injury.

Earlier version: (A) Heatmap displays the area under the curve (AUC) scores of transcription factor (TF) motifs, estimated per dot by Single-Cell Regulatory Network Inference and Clustering (SCENIC), highlighting gene regulatory networks and differentially activated TF motifs in hepatocyte zones (PP, Mid, PC) at 6 h post-APAP. Columns represent TF motifs, rows represent dots, and color intensity indicates AUC scores. (B) Heatmap shows the activity of transcription factors (TFs) across different zonal regions at 6 h post-APAP. TF activity was quantified using AUCell scoring. The color key represents regulon activity scores. (C) Violin plots shows transcriptional activity of Atf4 and Ddit3 in each zonation at 6 h post-APAP. (D) GO pathway analysis reveals the top 15 enriched pathways for genes co-regulated by Atf4 and Ddit3 in mid hepatocyte zones at 6 h post-APAP. (E) Heatmap showing the target genes expression of Atf4 (left) and Ddit3 (right) in hepatocyte zones (zones PP, Mid, and PC) at 6h post-APAP. (F) Immunohistochemistry staining of p-eIF2A in liver sections from mice at 0 h and 3 h post-APAP. p-eIF2A-positive hepatocytes are indicated by red arrows. The number of peIF2A-positive cells per field of view (FOV) is quantified. n=3 mice/group. (G) Heatmap displays stress response and ER stress-related genes from differentially expressed genes (DEGs) identified in Mid hepatocyte zones at 0 and 6 h post-APAP. Data are represented as means ± SD; *p < 0.05; **p < 0.01; ***p < 0.001; ****p < 0.0001; ns, not significant.

Revised version: (A) Heatmap displays the area under the curve (AUC) scores of transcription factor (TF) motifs, estimated per dot by Single-Cell Regulatory Network Inference and Clustering (SCENIC), highlighting gene regulatory networks and differentially activated TF motifs in hepatocyte zones (PP, Mid, PC) at 6 h post-APAP. Columns represent TF motifs, rows represent dots, and color intensity indicates AUC scores. (B) Heatmap shows inferred transcription factors (TFs) activity across different zonal regions at 6 h post-APAP. Activity was quantified as regulon enrichment scores using AUCell. The color scale represents regulon activity scores. (C) Violin plots shows regulon activity of Atf4 and Ddit3 in each zonation at 6 h post-APAP. Activity was quantified as regulon enrichment scores using AUCell. (D) GO pathway analysis reveals the top 15 enriched pathways for genes co-regulated by Atf4 and Ddit3 in mid hepatocyte zones at 6 h post-APAP. (E) Heatmap showing the target genes expression of Atf4 (left) and Ddit3 (right) in hepatocyte zones (zones PP, Mid, and PC) at 6h post-APAP. (F) Immunohistochemistry staining of p-eIF2A in liver sections from mice at 0 h and 3 h post-APAP. p-eIF2A-positive hepatocytes are indicated by red arrows. The number of p-eIF2A-positive cells per field of view (FOV) is quantified. n=3 mice/group. (G) Heatmap displays stress response and ER stress-related genes from differentially expressed genes (DEGs) identified in Mid hepatocyte zones at 0 and 6 h post-APAP. Data are represented as means ± SD; Unpaired Student’s t-test (F). *p < 0.05; **p < 0.01; ***p < 0.001; ****p < 0.0001; ns, not significant.

Earlier version: Among the top 10 TFs with highest activity in the Mid-zone during early AILI, Ddit3 and Atf3 were the two most highly expressed, with Atf4 ranking seventh (Figure 4B and Figure S3B). Since Ddit3 expression is often regulated by Atf4, and the Atf4-Ddit3 axis plays a critical role in stress adaptation,28 we chose to focus on this pathway for further investigation. The violin plot analysis demonstrated that Atf4 and Ddit3 were prominently active in Mid-zone, particularly at 3 and 6 hours post-APAP (Figure 4C and Figure S3C). Gene regulatory networks illustrated associations between Atf4 and its target genes (such as Gdf15, Hspa8, and Ddit3) and Ddit3 with its targets (such as Hspa8, Hmox1, Osgin1, Gadd45a, and Egr1) at 3 hours post-APAP (Figure S3D). Heatmap further clearly showed that the target genes of Atf4 and Ddit3 are exclusively expressed in the Mid-zone hepatocytes at 6 hours post-APAP (Figure 4D). Many of these target genes were DEGs in Mid-zone during early AILI, underscoring the critical role of the Atf4-Ddit3 axis in orchestrating gene expression changes in this region.

Revised version: Among the top 10 TFs with highest activity in the Mid-zone during early APAP-induced liver injury, Ddit3 and Atf3 were the two most highly expressed, with Atf4 ranking seventh (Figure 4B and Figure S3B). Since Ddit3 expression is often regulated by Atf4, and the Atf4-Ddit3 axis plays a critical role in stress adaptation (Zhang et al., 2022), we chose to focus on this pathway for further investigation. The violin plot analysis demonstrated that Atf4 and Ddit3 were prominently active in Mid-zone, particularly at 3 and 6 hours post-APAP (Figure 4C and Figure S3C). Gene regulatory networks illustrated associations between Atf4 and its target genes (such as Gdf15, Hspa8, and Ddit3) and Ddit3 with its targets (such as Hspa8, Hmox1, Osgin1, Gadd45a, and Egr1) at 3 hours post-APAP (Figure S3D). Heatmap further clearly showed that the target genes of Atf4 and Ddit3 are exclusively expressed in the Mid-zone hepatocytes at 6 hours post-APAP (Figure 4D). Many of these target genes were DEGs in Mid-zone during early APAP-induced liver injury, underscoring the critical role of the Atf4-Ddit3 axis in orchestrating gene expression changes in this region.

Earlier version: From the regulatory networks, it is evident that many genes regulated by Atf4 and Chop overlap. Therefore, we conducted a GO pathway analysis to identify the pathways co-regulated by Atf4 and Chop. The results revealed the top 15 enriched pathways for genes co-regulated by Atf4 and Chop in Mid-zone at 3 and 6 hours post-APAP (Figure 4E and Figure S3E). These pathways encompass responses to unfolded proteins, stress, autophagy, and regulation of apoptotic signaling pathways. Additionally, GO pathway analysis identified the top 15 enriched pathways for up-regulated genes specifically in Mid-zone at 3 hours post-APAP compared to other zones, further highlighting pathways co-regulated by Atf4 and Chop in response to APAP (Figure S3F). The pathways co-regulated by Atf4 and Chop point to the activation of ISR. To confirm the occurrence of ISR in Mid-zone during the initiation stage of AILI, we performed immunohistochemical staining for phosphorylated eIF2α (p-eIF2α) at 0 and 3 hours post-APAP and observed increased expression of p-eIF2α in the mid-zone (Figure 4F). While the ISR is a generalized pathway integrating endoplasmic reticulum (ER) stress and other insults through eIF2α kinases, ER stress is a subset of cellular stresses that activates the unfolded protein response (UPR) via PERK/IRE1/ATF6. Their overlap occurs at PERK-eIF2α-ATF4, but ER stress uniquely engages IRE1/XBP1 and ATF6 for ER-specific repair.15 To evaluate the involvement of ER stress in Mid-zone, we further analyzed the expression levels of genes related to stress response and ER stress based on DEGs identified in Mid-zone at 0, 3, and 6 hours post-APAP treatment (Figure 4G and Figure S3G). Most of stress response-related genes are highly upregulated (Figure 4G and Figure S3G). However, despite the high expression of Atf4 and Ddit3, other ER stress-related genes, such as Eif2ak3, Atf6, and Xbp1 28 either remained unchanged or were downregulated at 3 and 6 hours compared to 0 hour post-APAP (Figure 4G and Figure S3G). This analysis highlights the unique role of the Atf4-Ddit3 axis in Mid-zone during early AILI.

Revised version: From the regulatory networks, it is evident that many genes regulated by Atf4 and Chop overlap. Therefore, we conducted a GO pathway analysis to identify the pathways co-regulated by Atf4 and Chop. The results revealed the top 15 enriched pathways for genes co-regulated by Atf4 and Chop in Mid-zone at 3 and 6 hours post-APAP (Figure 4E and Figure S3E). These pathways encompass responses to unfolded proteins, stress, autophagy, and regulation of apoptotic signaling pathways. Additionally, GO pathway analysis identified the top 15 enriched pathways for up-regulated genes specifically in Mid-zone at 3 hours post-APAP compared to other zones, further highlighting pathways co-regulated by Atf4 and Chop in response to APAP (Figure S3F). The pathways co-regulated by Atf4 and Chop point to the activation of ISR. To confirm the occurrence of ISR in Mid-zone during the initiation stage of APAP-induced liver injury, we performed immunohistochemical staining for phosphorylated eIF2α (p-eIF2α) at 0 and 3 hours post-APAP and observed increased expression of p-eIF2α in the mid-zone (Figure 4F). While the ISR is a generalized pathway integrating endoplasmic reticulum (ER) stress and other insults through eIF2α kinases, ER stress is a subset of cellular stresses that activates the unfolded protein response (UPR) via Perk/Ire1/Atf6. Their overlap occurs at Perk-eIF2α-Atf4, but ER stress uniquely engages Ire1/Xbp1 and Atf6 for ER-specific repair (Costa-Mattioli and Walter, 2020). To evaluate the involvement of ER stress in Mid-zone, we further analyzed the expression levels of genes related to stress response and ER stress based on DEGs identified in Mid-zone at 0, 3, and 6 hours post-APAP treatment (Figure 4G and Figure S3G). Most of stress response-related genes are highly upregulated (Figure 4G and Figure S3G). However, despite the high expression of Atf4 and Ddit3, other ER stress-related genes, such as Eif2ak3, Atf6, and Xbp1 (Zhang et al., 2022) either remained unchanged or were downregulated at 3 and 6 hours compared to 0 hour post-APAP (Figure 4G and Figure S3G). This analysis highlights the unique role of the Atf4-Ddit3 axis in Mid-zone during early APAP-induced liver injury.

Earlier version: To validate Mid-zone responses and transcriptional changes during early AILI, we analyzed an additional scRNA-seq dataset from AILI mice.29 UMAP visualization demonstrated that hepatocytes could be clustered into three distinct zones (PP, PC, and Mid) at 3 and 6 hours post-APAP (Figures S4A and S4B). The zonal gene expression patterns were confirmed, with PC-specific genes (Glul, Cyp2e1, Cyp2a5), PP-specific genes (Alb, Cyp2f2, Sds), and Mid-zone genes (Hmox1, Atf3, Ubc) exhibiting distinct spatial distributions (Figure S4C). GO pathway analysis revealed the top 10 enriched pathways for upregulated genes in Mid-zonal hepatocytes at 3 and 6 hours post-APAP, highlighting key stress-response and cell death processes (Figures S4D and S4E). Consistent with our findings, SCENIC analysis also identified the activation of similar TF motifs across zonal regions, with distinct Ddit3 and Atf3 motifs activated in Mid-zone hepatocytes at 3 and 6 hours post-APAP (Figures S4F and S4G). These results reinforce the regulatory changes observed in our ST data, further emphasizing the unique role of mid-zonal hepatocytes in stress adaptation.

Revised version: To validate Mid-zone responses and transcriptional changes during early APAP-induced liver injury, we analyzed an additional scRNA-seq dataset from APAP-induced liver injury mice (Matchett et al., 2024). UMAP visualization demonstrated that hepatocytes could be clustered into three distinct zones (PP, PC, and Mid) at 3 and 6 hours post-APAP (Figures S4A and S4B). The zonal gene expression patterns were confirmed, with PC-specific genes (Glul, Cyp2e1, Cyp2a5), PP-specific genes (Alb, Cyp2f2, Sds), and Mid-zone genes (Hmox1, Atf3, Ubc) exhibiting distinct spatial distributions (Figure S4C). GO pathway analysis revealed the top 10 enriched pathways for upregulated genes in Mid-zonal hepatocytes at 3 and 6 hours post-APAP, highlighting key stress-response and cell death processes (Figures S4D and S4E). Consistent with our findings, SCENIC analysis also identified the activation of similar TF motifs across zonal regions, with distinct Ddit3 and Atf3 motifs activated in Mid-zone hepatocytes at 3 and 6 hours post-APAP (Figures S4F and S4G). These results reinforce the regulatory changes observed in our ST data, further emphasizing the unique role of mid-zonal hepatocytes in stress adaptation.

The Atf4-Ddit3 axis protects hepatocytes from liver injury

Earlier version: To investigate the role of Atf4-Ddit3 axis in cell fate determination, we first examined Atf4 expression in liver sections by immunohistochemical staining at 0 and 6 hours post-APAP. The results showed gradual increased expression of Atf4 in Mid-zone during early AILI (Figure 5A). Subsequently, we employed the adeno-associated virus 8 (AAV8) and thyroxine-binding globulin (TBG) promoter method to overexpress Atf4 specifically in hepatocytes (Figure 5B). Atf4 overexpression significantly reduced AILI, as evidenced by H&E staining and serum ALT levels (Figure 5C and 5D). Furthermore, immunohistochemical staining of Atf4 indicated its localization in hepatocyte nuclei around the injured area (Figure 5E). Immunofluorescence and immunohistochemical staining for Ki67 and p-eIF2α, respectively, showed that Atf4 overexpression notably decreased Ki67 expression and increased p-eIF2α expression (Figure 5F and 5G). Together, these findings suggest that elevated Atf4 expression enhances ISR, offering stress protection under conditions of reduced proliferation.

Revised version: To investigate the role of Atf4-Ddit3 axis in cell fate determination, we first examined Atf4 expression in liver sections by immunohistochemical staining at 0 and 6 hours post-APAP. The results showed gradual increased expression of Atf4 in Mid-zone during early APAP-induced liver injury (Figure 5A). Subsequently, we employed the adeno-associated virus 8 (AAV8) and thyroxine-binding globulin (TBG) promoter method to overexpress Atf4 specifically in hepatocytes (Figure 5B). Atf4 overexpression significantly reduced APAP-induced liver injury, as evidenced by H&E staining and serum ALT levels (Figure 5C and 5D). Furthermore, immunohistochemical staining of Atf4 indicated its localization in hepatocyte nuclei around the injured area (Figure 5E). Immunofluorescence and immunohistochemical staining for Ki67 and p-eIF2α, respectively, showed that Atf4 overexpression notably decreased Ki67 expression and increased p-eIF2α expression (Figure 5F and 5G). Together, these findings suggest that elevated Atf4 expression enhances ISR, offering stress protection under conditions of reduced proliferation.

The Atf4-Ddit3 axis protects hepatocytes from liver injury.

Earlier version: (A) Immunohistochemical staining of Atf4 in liver sections at 0 and 6 h post-APAP. Red arrows indicate Atf4-positive hepatocytes. Zonal distribution of Atf4-positive cells in liver sections at 6 h post-APAP is quantified. The statistic is the percentage of Atf4-positive hepatocytes in each layer over the total number of Atf4-positive hepatocytes. n=3 mice. (B) Schematic figure illustrating the overexpression of Atf4 via AAV in hepatocytes of wildtype C57BL/6J mice. Overexpression of EGFP is used as a control. (C) H&E staining showing liver morphology from mice overexpressing AAV-TBG-EGFP or AAV-TBG-Atf4 at 6 h post-APAP. The injured area is outlined by black dashed lines. The percentage of injured area is quantified. n=3 mice/group. (D) Serum levels of ALT is measured in mice overexpressing AAV-TBG-EGFP or AAV-TBGAtf4 at 6, 24, 48, and 72 h post-APAP. n=6 mice/group. (E) Immunohistochemistry staining of Atf4 in liver sections from mice overexpressing AAVTBG-EGFP or AAV-TBG-Atf4 at 12 h post-APAP. Atf4-positive hepatocytes are indicated by red arrows. The number of Atf4-positive cells per field of view (FOV) is quantified. n=9 mice/group. (F) Immunofluorescence staining of Ki67 (red) was performed to assess proliferating hepatocytes in mice overexpressing AAV-TBG-EGFP or AAV-TBG-Atf4 at 72 h post-APAP. Cell nuclei were stained with DAPI (blue). The number of Ki67-positive cells per FOV is quantified. n=6 mice/group. (G) Immunohistochemistry staining of p-eIF2A in liver sections from mice overexpressing AAVTBG-EGFP or AAV-TBG-Atf4 at 12 h post-APAP. p-eIF2A-positive hepatocytes are indicated by red arrows. The number of Atf4-positive cells per FOV is quantified. n=6 mice/group. Data are represented as means ± SD; *p < 0.05; **p < 0.01; ***p < 0.001; ****p < 0.0001; ns, not significant.

Revised version: (A) Immunohistochemical staining of Atf4 in liver sections at 0 and 6 h post-APAP. Red arrows indicate Atf4-positive hepatocytes. Zonal distribution of Atf4-positive cells in liver sections at 6 h post-APAP is quantified. The statistic is the percentage of Atf4-positive hepatocytes in each layer over the total number of Atf4-positive hepatocytes. n=3 mice. (B) Schematic figure illustrating the overexpression of Atf4 via AAV in hepatocytes of wildtype C57BL/6J mice. Overexpression of EGFP is used as a control. (C) H&E staining showing liver morphology from mice overexpressing AAV-TBG-EGFP or AAV-TBG-Atf4 at 6 h post-APAP. The injured area is outlined by black dashed lines. The percentage of injured area is quantified. n=3 mice/group. Necrotic areas outlined by loss of cellular architecture on H&E. (D) Serum levels of ALT is measured in mice overexpressing AAV-TBG-EGFP or AAV-TBG-Atf4 at 6, 24, 48, and 72 h post-APAP. n=6 mice/group. (E) Immunohistochemistry staining of Atf4 in liver sections from mice overexpressing AAV-TBG-EGFP or AAV-TBG- Atf4 at 24 h post-APAP. Atf4-positive hepatocytes are indicated by red arrows. The number of Atf4-positive cells per field of view (FOV) is quantified. n=9 mice/group. (F) Immunofluorescence staining of Ki67 (red) was performed to assess proliferating hepatocytes in mice overexpressing AAV-TBG-EGFP or AAV-TBG-Atf4 at 72 h post-APAP. Cell nuclei were stained with DAPI (blue). The number of Ki67-positive cells per FOV is quantified. n=6 mice/group. (G) Immunohistochemistry staining of p-eIF2A in liver sections from mice overexpressing AAV-TBG-EGFP or AAV-TBG- Atf4 at 24 h post-APAP. p-eIF2A-positive hepatocytes are indicated by red arrows. The number of p-eIF2A-positive cells per FOV is quantified. n=6 mice/group. Data are represented as means ± SD; Unpaired Student’s t-test (C-G). One-way ANOVA (A). *p < 0.05; **p < 0.01; ***p < 0.001; ****p < 0.0001; ns, not significant.

The Atf4-Ddit3 axis protects the liver by pausing proliferation via Btg2

Earlier version: Atf4 often works with Ddit3 to mediate stress response, particularly in the context of the ISR.23,24,30 Given the low expression of Atf4 and high expression of Ddit3 in the Mid-zone, we performed Cut&Run to identify the genomic regions where Ddit3 binds during the initiation phases of AILI. This approach was used to investigate how the Atf4-Ddit3 axis regulates the proliferation of Mid-zone hepatocytes. Compared with the input control, we observed heightened Ddit3-binding signals within the chromatin domains (Figure 6A). We further overlapped 6157 genes where Ddit3 binds, as identified by Cut&Run, with 47 DEGs specifically in the Mid-zone hepatocytes, identifying 39 overlapping genes termed Ddit3-binding DEGs (Figure 6B). Integrative Genomics Viewer (IGV) plots illustrate Ddit3-binding peaks and input controls at specific genomic sites such as Btg2, Atf3, Dnajb1, and Egr1 regulatory elements (Figure 6C).

Revised version: Atf4 often works with Ddit3 to mediate stress response, particularly in the context of the ISR (Kaspar et al., 2021; Ohoka et al., 2005; Woo et al., 2009). Given the low expression of Atf4 and high expression of Ddit3 in the Mid-zone, we performed Cut&Run to identify the genomic regions where Ddit3 binds during the initiation phases of APAP-induced liver injury. This approach was used to investigate how the Atf4-Ddit3 axis regulates the proliferation of Mid-zone hepatocytes. Compared with the input control, we observed heightened Ddit3-binding signals within the chromatin domains (Figure 6A). We further overlapped 6157 genes where Ddit3 binds, as identified by Cut&Run, with 47 DEGs specifically in the Mid-zone hepatocytes, identifying 39 overlapping genes termed Ddit3-binding DEGs (Figure 6B). Integrative Genomics Viewer (IGV) plots illustrate Ddit3-binding peaks and input controls at specific genomic sites such as Btg2, Atf3, Dnajb1, and Egr1 regulatory elements (Figure 6C).

The Atf4-Ddit3 axis protects the liver by pausing proliferation via Btg2.

Earlier version: (A) Heatmap of Ddit3-binding signals in chromatin domains. (B) Overlap of Ddit3-binding genes by Cut&Run with differentially expressed genes (DEGs) in mid zone was shown. These genes were referred as Ddit3-binding DEGs. (C) Integrative Genomics Viewer (IGV) plot shows the Cut&Run peaks of Ddit3 and input on Btg2-site, Ifrd1-site, Fos-site and Egr1-site. (D) Immunohistochemical staining of Btg2 in liver sections at 0 and 6 h post-APAP. Red arrows indicate Btg2-positive hepatocytes. (E) Schematic figure illustrating overexpression of Ddit3 via AAV in hepatocytes of wildtype C57BL/6J mice. Overexpression of EGFP is used as a control. (F) Immunohistochemical staining of Btg2 in liver sections from mice overexpressing either AAV-TBG-EGFP or AAV-TBG-Ddit3 at 12 h post-APAP. Btg2-positive hepatocytes are indicated by red arrows. Btg2-positive hepatocytes in liver sections is quantified. n=4 mice/group. (G) Schematic figure illustrating overexpression of Btg2 via AAV in hepatocytes of wildtype C57BL/6J mice. Overexpression of EGFP is used as a control. (H) H&E staining shows morphology of livers from AAV-TBG-EGFP/Btg2 overexpressing mice at 6 h post-APAP. Injured area is outlined by black dashed lines. The percentage of injury area is quantified. n=3 mice/group. (I) Serum levels of ALT is measured at 6h post-APAP in AAV-TBG-EGFP/Btg2 overexpressing mice. n=3 mice/group. (J) Immunohistochemical staining of Ki67 in liver sections from mice overexpressing either AAV-TBG-EGFP or AAV-TBG-Btg2 at 72 h post-APAP. Ki67-positive hepatocytes are indicated by red arrows. Ki67-positive cells in liver sections is quantified. n=4 mice/group. (K) Immunohistochemical staining of Ki67 in liver sections from mice after knocking down Btg2 via overexpressing either IP479-EFS-sgEGFP or IP479-EFS-sgBtg2 at 72 h post-APAP. Ki67-positive hepatocytes are indicated by red arrows. Ki67-positive cells in liver sections is quantified. n=4 mice/group. Data are represented as means ± SD; **p < 0.01; ***p < 0.001; ****p < 0.0001; ns, not significant.

Revised version: (A) Heatmap of Ddit3-binding signals in chromatin domains. (B) Overlap of Ddit3-binding genes by Cut&Run with differentially expressed genes (DEGs) in mid zone was shown. These genes were referred as Ddit3-binding DEGs. (C) Integrative Genomics Viewer (IGV) plot shows the Cut&Run peaks of Ddit3 and input on Btg2-site, Ifrd1-site, Fos-site and Egr1-site. (D) Immunohistochemical staining of Btg2 in liver sections at 0 and 6 h post-APAP. Red arrows indicate Btg2-positive hepatocytes. (E) Schematic figure illustrating overexpression of Ddit3 via AAV in hepatocytes of wildtype C57BL/6J mice. Overexpression of EGFP is used as a control. (F) Immunohistochemical staining of Btg2 in liver sections from mice overexpressing either AAV-TBG-EGFP or AAV-TBG-Ddit3 at 12 h post-APAP. Btg2-positive hepatocytes are indicated by red arrows. Btg2-positive hepatocytes in liver sections is quantified. n=4 mice/group. (G) Schematic figure illustrating overexpression of Btg2 via AAV in hepatocytes of wildtype C57BL/6J mice. Overexpression of EGFP is used as a control. (H) H&E staining shows morphology of livers from AAV-TBG-EGFP or AAV-TBG-Btg2 overexpressing mice at 6 h post-APAP. Injured area is outlined by black dashed lines. The percentage of injury area is quantified. n=3 mice/group. Necrotic areas outlined by loss of cellular architecture on H&E. (I) Serum levels of ALT is measured at 6h post-APAP in AAV-TBG-EGFP or AAV-TBG-Btg2 overexpressing mice. n=3 mice/group. (J) Immunohistochemical staining of Ki67 in liver sections from mice overexpressing either AAV-TBG-EGFP or AAV-TBG-Btg2 at 72 h post-APAP. Ki67-positive hepatocytes are indicated by red arrows. Ki67-positive cells in liver sections is quantified. n=4 mice/group. (K) Schematic figure illustrating knockdown of Btg2 via AAV in hepatocytes of wildtype C57BL/6J mice. IP479-EFS-sgEGFP is used as a control. (L) Serum levels of ALT is measured at 6h post-APAP in IP479-EFS-sgEGFP or IP479-EFS-sgBtg2 knockdown mice. n=3 mice/group. (M) H&E staining shows morphology of livers from IP479-EFS-sgEGFP or IP479-EFS-sgBtg2 knockdown mice at 6 h post-APAP. Injured area is outlined by black dashed lines. The percentage of injury area is quantified. n=3 mice/group. Necrotic areas outlined by loss of cellular architecture on H&E. (N) Immunohistochemical staining of Ki67 in liver sections from IP479-EFS-sgEGFP or IP479-EFS-sgBtg2 knockdown mice at 72 h post-APAP. Ki67-positive cells in liver sections is quantified. n=3 mice/group. Data are represented as means ± SD; Unpaired Student’s t-test (F, H-N). **p < 0.01; ***p < 0.001; ****p < 0.0001; ns, not significant.

Earlier version: Among the Ddit3-binding DEGs, we observed B-cell translocation gene 2 (Btg2), which is cell cycle inhibitor.31–33 Moreover, Btg2 is one of the target genes that coregulated by Ddit3 and Atf4. Immunohistochemical staining of Btg2 in liver sections at 0 and 6 hours post-APAP revealed increased expression of Btg2 around the injured area in response to APAP (Figure 6D). Additionally, we investigated whether Ddit3 could upregulate Btg2 expression in Mid-zone. We overexpressed Ddit3 in the hepatocytes of wildtype C57BL/6J mice using the AAV8 method, with EGFP overexpression serving as a control (Figure 6E). Immunohistochemical staining of Btg2 in liver sections from mice overexpressing either AAV-TBG-EGFP or AAV-TBG-Ddit3 at 12 hours post-APAP indicated that Ddit3 overexpression significantly upregulated Btg2 expression post-APAP (Figure 6F). This result suggests that Btg2 is a downstream target of Ddit3.

Revised version: Among the Ddit3-binding DEGs, we observed B-cell translocation gene 2 (Btg2), which is cell cycle inhibitor (Park et al., 2008; Stupfler et al., 2016; Yuniati et al., 2019). Moreover, Btg2 is one of the target genes that coregulated by Ddit3 and Atf4. Immunohistochemical staining of Btg2 in liver sections at 0 and 6 hours post-APAP revealed increased expression of Btg2 around the injured area in response to APAP (Figure 6D). Additionally, we investigated whether Ddit3 could upregulate Btg2 expression in Mid-zone. We overexpressed Ddit3 in the hepatocytes of wildtype C57BL/6J mice using the AAV8 method, with EGFP overexpression serving as a control (Figure 6E). Immunohistochemical staining of Btg2 in liver sections from mice overexpressing either AAV-TBG-EGFP or AAV-TBG-Ddit3 at 12 hours post-APAP indicated that Ddit3 overexpression significantly upregulated Btg2 expression post-APAP (Figure 6F). This result suggests that Btg2 is a downstream target of Ddit3.

Added in the revised version: To explore the role of Btg2 in liver injury and hepatocyte proliferation, we overexpressed Btg2 in hepatocytes of wild-type C57BL/6J mice (Figure 6G). H&E staining and serum ALT levels indicated that Btg2 overexpression significantly reduced APAP-induced liver injury (Figure 6H and 6I). Ki67 staining further demonstrated that Btg2 overexpression significantly decreased hepatocyte proliferation (Figure 6J). To assess potential confounds from AAV-based overexpression, we performed western blot analysis of Cyp2e1 protein levels in AAV-EGFP control mice and mice overexpressing Atf4 or Btg2 without APAP treatment. Compared with the AAV-EGFP group, Cyp2e1 expression was slightly decreased in the AAV-Atf4 and AAV-Btg2 groups (Figure S5A). Importantly, this reduction is well below the threshold of Cyp inhibition known to confer significant protection against APAP-induced liver injury (Ganetsky et al., 2019; Ye et al., 2022a). Therefore, the modest decrease in Cyp2e1 expression observed here is unlikely to account for the pronounced protective effects of Atf4 or Btg2 overexpression. We also quantified the penetrance of AAV-mediated overexpression. Mice were injected with 1.2 × 10¹¹ viral genome copies per animal, yielding average transduction efficiencies of 32%, 18%, and 23% for AAV-EGFP, AAV-Atf4, and AAV-Btg2, respectively (Figure S5B). Individual animal transduction efficiency showed a negative correlation with serum ALT levels (e.g. Pearson r = –0.7681, p = 0.0260 for Atf4; Figure S5C), demonstrating that greater transgene expression associates with stronger protection. Moreover, we used the CasRx-sgRNA system (Zhou et al., 2020) to knock down Btg2 by overexpressing either IP479-EFS-sgEGFP or IP479-EFS-sgBtg2 (Figure 6K). Knockdown of Btg2 using AAV8- CasRx achieved a moderate (∼30%) reduction in Btg2 expression (Figure S5D). Despite this partial efficiency, we observed a significant increase in serum ALT levels (∼2- fold) (Figure 6L), expansion of necrotic areas (∼1.5- fold) (Figure 6M), and a marked increase in Ki67⁺ hepatocytes (∼1.8- fold) compared to control mice (Figure 6N). Together, these findings illustrate Atf4-Ddit3 axis inhibits hepatocytes proliferation in Mid-zone via Btg2 during early APAP-induced liver injury.

Added in the revised version: To test whether the proposed mechanism applies to other liver injuries, we employed mouse models of partial hepatectomy (PHx) and carbon tetrachloride (CCl4)-induced acute liver injury. CCl4 causes centrilobular necrosis (similar to acetaminophen) but induces damage via free radical–mediated mechanisms rather than through cytochrome P450–dependent reactive metabolite formation (Wang et al., 2024). In our CCl4 model (administered intraperitoneally in corn oil, with samples collected 18 h post- injection), the ISR was activated around injury sites, accompanied by decreased proliferation, as evidenced by increased expression of p- eIF2α, Atf4, Chop, and Btg2, along with reduced Ki67 expression (Figure S6A–G). Partial hepatectomy involves mechanical resection, resulting in ischemic injury in the remaining portion of the left lobe due to compromised blood supply (Qu et al., 2022). In this model (examined 24 h after surgery), ISR activation was similarly observed around ischemic injury sites, with increased expression of p- eIF2α, Atf4, Chop, and Btg2 and undetectable Ki67 expression (Figure S7A–G). Together, these findings suggest that the proposed mechanism may be broadly applicable to other types of liver injury.

Removed in the revised version: To explore the role of Btg2 in liver injury and hepatocyte proliferation, we further overexpressed Btg2 in hepatocytes of wildtype C57BL/6J mice (Figure 6G). H&E staining and serum ALT levels indicated that Btg2 overexpression significantly reduced AILI (Figure 6H and 6I). Ki67 staining further demonstrated that Btg2 overexpression significantly decreased hepatocyte proliferation (Figure 6J). Moreover, we used the CasRx-sgRNA system34 to knock down Btg2 by overexpressing either IP479-EFS-sgEGFP or IP479-EFS-sgBtg2. In contrast to Btg2 overexpression, Btg2 knockdown resulted in increased Ki67 expression in hepatocytes, suggesting enhanced hepatocyte proliferation (Figure 6K). Together, these findings illustrate Atf4-Ddit3 axis inhibits hepatocytes proliferation in Mid-zone via Btg2 during early AILI.

Discussion

Earlier version: Our study reveals that during the early phase of AILI, mid-zone hepatocytes undergo a transient proliferation arrest, orchestrated by a distinctive transcriptional program centered on the Atf4-Chop axis. Spatial transcriptomics and functional assays demonstrate that this adaptive response promotes survival by upregulating the cell cycle inhibitor Btg2, highlighting a trade-off between cytoprotection and proliferative capacity. Notably, zonal metabolic heterogeneity in APAP processing drives preferential activation of these stress pathways in mid-zone hepatocytes. These findings uncover a critical mechanism by which hepatocytes prioritize stress adaptation before initiating regeneration, offering new insights into the spatiotemporal regulation of liver repair (Figure 7).

Revised version: Our study reveals that during the early phase of APAP-induced liver injury, mid-zone hepatocytes undergo a transient proliferation arrest, orchestrated by a distinctive transcriptional program centered on the Atf4-Chop axis. Spatial transcriptomics and functional assays demonstrate that this adaptive response promotes survival by upregulating the cell cycle inhibitor Btg2, highlighting a trade-off between cytoprotection and proliferative capacity. Notably, zonal metabolic heterogeneity in APAP processing drives preferential activation of these stress pathways in mid-zone hepatocytes. These findings uncover a critical mechanism by which hepatocytes prioritize stress adaptation before initiating regeneration, offering new insights into the spatiotemporal regulation of liver repair (Figure 7).

The Atf4-Ddit3 axis mediates integrated stress protection in mid-zone hepatocytes at the expense of proliferative capacity during early AILI.

Schematic model depicting the adaptive response of mid-zone hepatocytes to acute liver injury, wherein activation of the Atf4-Ddit3 pathway upregulates Btg2 to induce transient cell cycle arrest, balancing survival and regeneration.

Earlier version: Our study reveals that mid-zone hepatocytes employ a protective proliferation arrest during early AILI, mediated by the Atf4-Chop axis of the integrated stress response (ISR). While the ISR shares upstream initiators with ER stress (e.g., eIF2α phosphorylation), its outcomes are context-dependent. ER stress specifically activates the unfolded protein response (UPR) through sensors inositol-requiring enzyme 1α (IRE1α), activating transcription factor 6 (ATF6), and protein kinase RNA-like ER kinase (PERK),35 with genetic studies demonstrating their critical roles in AILI – knockout of IRE1α, X-box binding protein 1 (XBP1), or Chop consistently attenuates liver damage, confirming the UPR’s contribution to hepatotoxicity.19,36–38 Unlike classical ER stress, we find that mid-zone hepatocytes leverage the Atf4-Chop axis to transiently halt proliferation via Btg2, favoring survival. This aligns with emerging evidence that Chop can paradoxically support adaptation in mitochondrial stress,30,39 suggesting that its role in fine-turning of ISR in mammals.

Revised version: Our study reveals that mid-zone hepatocytes employ a protective proliferation arrest during early APAP-induced liver injury, mediated by the Atf4-Chop axis of the integrated stress response (ISR). While the ISR shares upstream initiators with ER stress (e.g., eIF2α phosphorylation), its outcomes are context-dependent. ER stress specifically activates the UPR through sensors inositol-requiring enzyme 1α (Ire1α), activating transcription factor 6 (Atf6), and protein kinase RNA-like ER kinase (Perk) (Wang et al., 2006), with genetic studies demonstrating their critical roles in APAP-induced liver injury knockout of Ire1α, X-box binding protein 1 (Xbp1), or Chop consistently attenuates liver damage, confirming the UPR’s contribution to hepatotoxicity (Hur et al., 2012; Nagy et al., 2007; Uzi et al., 2013; Ye et al., 2022b). Unlike classical ER stress, we find that mid-zone hepatocytes leverage the Atf4-Chop axis to transiently halt proliferation via Btg2, favoring survival. This aligns with emerging evidence that Chop can paradoxically support adaptation in mitochondrial stress (Kaspar et al., 2021; Kress et al., 2023), suggesting that its role in fine-turning of ISR in mammals.

Added in the revised version: Our findings on the early stress-induced proliferation arrest in mid-lobular hepatocytes (3–12h post-APAP) complement yet fundamentally extend prior knowledge of zonal liver regeneration. While foundational studies established perinecrotic hepatocytes as drivers of later proliferation (>24h) (Bajt et al., 2003), and recent work delineated CXCR2+ hepatocytes as pro-proliferative effectors at moderate doses (Nguyen et al., 2022) or CXCL14/p21-mediated senescence at high doses (Umbaugh et al., 2024), our work resolves an earlier, spatiotemporally distinct adaptive phase. Crucially, we identify mid-zone hepatocytes—not merely as transitional cells en route to proliferation—but as orchestrators of an active survival strategy via the Atf4-Chop-Btg2 axis, which redirects resources from division to cytoprotection (e.g., redox defense). This mechanism operates before necrosis expansion and is distinct from periportal/pericentral stress responses, which focus on 2h post-injury and lack mid-zone resolution (Umbaugh et al., 2021). By integrating spatial transcriptomics with functional validation (Cut&Run, AAV-Btg2 modulation), we establish causal links between stress signaling and quiescence—unlike descriptive profiling in earlier zonal studies (Wu et al., 2024). Thus, our model bridges early stress responses and later regeneration, revealing how mid-zone hepatocytes enable subsequent repair by prioritizing survival during the critical initiation window.

Earlier version: The antiproliferative activity of Btg2 has been studied through an integrated network of transcriptional and post-transcriptional cell cycle regulatory mechanisms. A study reveals that Btg2 functions as a molecular bridge between poly(A)-binding protein cytoplasmic 1 (PABPC1) and the CCR4-NOT deadenylase complex subunit CNOT7 (CAF1), facilitating accelerated degradation of cell cycle-related transcripts through enhanced poly(A) tail removal.31 At the transcriptional level, Btg2 exerts specific control over G2/M progression by disrupting the positive feedback loop between cyclin B1 and its transcriptional activator FOXM1 (forkhead box protein M1), thereby suppressing cyclin B1 expression and inducing cell cycle arrest.32 Furthermore, Btg2 demonstrates broader cell cycle inhibitory capacity through its interactions with cyclin-dependent kinases (CDKs) during G1/S transition and its ability to modulate the tumor suppressor p53 (TP53).33 In the context of AILI, the stress-responsive Atf4-Chop pathway induces Btg2 expression in mid-zonal hepatocytes, where it may coordinate a comprehensive proliferation arrest through these complementary mechanisms. This multilayered regulatory strategy enables hepatocytes to temporarily suspend cell division and redirect resources toward cellular repair and stress adaptation during the acute phase of toxic injury.

Revised version: The antiproliferative activity of Btg2 has been studied through an integrated network of transcriptional and post-transcriptional cell cycle regulatory mechanisms. A study reveals that Btg2 functions as a molecular bridge between poly(A)-binding protein cytoplasmic 1 (Pabpc1) and the CCR4-NOT deadenylase complex subunit CNOT7 (Caf1), facilitating accelerated degradation of cell cycle-related transcripts through enhanced poly(A) tail removal (Stupfler et al., 2016). At the transcriptional level, Btg2 exerts specific control over G2/M progression by disrupting the positive feedback loop between cyclin B1 and its transcriptional activator Foxm1 (forkhead box protein M1), thereby suppressing cyclin B1 expression and inducing cell cycle arrest (Park et al., 2008). Furthermore, Btg2 demonstrates broader cell cycle inhibitory capacity through its interactions with cyclin-dependent kinases (Cdks) during G1/S transition and its ability to modulate the tumor suppressor p53 (TP53) (Yuniati et al., 2019). In the context of APAP-induced liver injury, the stress-responsive Atf4-Chop pathway induces Btg2 expression in mid-zonal hepatocytes, where it may coordinate a comprehensive proliferation arrest through these complementary mechanisms. This multilayered regulatory strategy enables hepatocytes to temporarily suspend cell division and redirect resources toward cellular repair and stress adaptation during the acute phase of toxic injury.

Earlier version: A key limitation of this study is the scarcity of early-phase AILI human samples, which currently precludes definitive validation of the clinical relevance of these zonal adaptations. To address this gap, we performed a preliminary analysis of published single-cell RNA sequencing data from APAP-overdose patients.29 Intriguingly, in one of two analyzed patients, mid-zone hepatocytes exhibited transcriptional signatures remarkably consistent with our murine findings, including: (1) upregulation of Atf4-Chop pathways, and (2) downregulation of cell proliferation genes. While the small sample size (n=2) prevents definitive conclusions, these preliminary observations suggest potential conservation of stress-responsive mechanisms between murine and human mid-zone hepatocytes during AILI. Further studies with larger clinical cohorts will be essential to verify these findings and establish their translational significance.

Revised version: A key limitation of this study is the scarcity of early-phase APAP-induced liver injury human samples, which currently precludes definitive validation of the clinical relevance of these zonal adaptations. To address this gap, we performed a preliminary analysis of published ST data from APAP-overdose patients (Matchett et al., 2024). Intriguingly, in one of two analyzed patients, mid-zone hepatocytes exhibited transcriptional signatures remarkably consistent with our murine findings, including: (1) upregulation of Atf4-Chop pathways, and (2) downregulation of cell proliferation genes. While the small sample size (n=2) prevents definitive conclusions, these preliminary observations suggest potential conservation of stress-responsive mechanisms between murine and human mid-zone hepatocytes during APAP-induced liver injury. Further studies with larger clinical cohorts will be essential to verify these findings and establish their translational significance. While our ST analysis and Ki67 staining consistently indicated hepatocytes as the primary proliferating compartment during early APAP-induced liver injury, we recognize that non-parenchymal cells (e.g., immune cells, stellate cells) may contribute to division signals. Single-cell resolution techniques would be required to completely exclude these contributions, representing a technical limitation of bulk spatial transcriptomics. Additionally, our ST data alone do not independently confirm zonal cell cycle arrest: ST lacks overt separation in S/G2/M scores, uses non-canonical markers (e.g., Nasp, Cks1b), and has low sensitivity for rare Ki67⁺ events (∼1% of spots). Thus, our primary evidence for midlobular proliferation dynamics comes from Ki67 immunohistochemistry (Figure 1), with ST findings (Figure 2E) shown as exploratory and directionally consistent with transient impairment, but not as definitive proof of cell cycle arrest.

Earlier version: In summary, our study uncovers a spatially coordinated stress adaptation mechanism in mid-zone hepatocytes during early AILI, where the Atf4-Chop axis transiently suppresses proliferation via Btg2 upregulation. By elucidating how hepatocytes prioritize survival before initiating repair, our work provides a framework for understanding the dynamic interplay between stress adaptation and regeneration in tissues.

Revised version: In summary, our study uncovers a spatially coordinated stress adaptation mechanism in mid-zone hepatocytes during early APAP-induced liver injury, where the Atf4-Chop axis transiently suppresses proliferation via Btg2 upregulation. By elucidating how hepatocytes prioritize survival before initiating repair, our work provides a framework for understanding the dynamic interplay between stress adaptation and regeneration in tissues.

Earlier version: Methods and materials

Revised version: Materials and methods

Added in the revised version: Animal experiments and procedures

Added in the revised version: Animals

Added in the revised version: C57BL/6J mice (strain no. N000013) were used as wild-type (WT) mice. All mouse colonies were maintained at the Animal Core Facility of Yunnan University. The animal studies were approved by the Yunnan University Institutional Animal Care and Use Committee (IACUC, Approval No. YNU20220387). For APAP treatment, mice (8-12 weeks old) were fasted overnight (5:00pm to 9:00am) before i.p. injected with APAP (Sigma, A7085) at a dose of 300 mg/kg for male mice, as female mice are less susceptible to APAP-induced liver injury (Guerrero Munoz and Fearon, 1984). Male mice have been the choice in the vast majority of the studies of APAP-induced liver injury reported in the literature (Ben-Moshe et al., 2022b; Campion et al., 2008; Matchett et al., 2024; Shan et al., 2021). Therefore, we used male mice in the majority of the experiments presented. In one experiment, mice were i.p. injected with various doses of APAP (100, 300, 500mg/kg). For carbon tetrachloride (CCl4) treatment, mice were i.p. injected with corn oil or CCl4 (Macklin, Cat #C805329) at a dose of 0.8ul/g (corn oil dilute to 20%). For 70% partial hepatectomy, mice were placed in a supine position and secured under isoflurane inhalation anesthesia. A midline upper abdominal incision is made. The liver is gently exposed, and the left lateral lobe and the median lobe (comprising the left and right median lobes) are mobilized. The hepatic pedicle at the base of each lobe is ligated with 6-0 silk suture. After confirming the absence of bleeding, the lobes are excised just distal to the ligature, removing approximately 70% of the liver mass. The resection site is inspected for active bleeding. The abdominal incision is closed in layers with absorbable suture. Postoperatively, the mouse is placed on a heating pad for recovery.

Added in the revised version: In some experiments, APAP-treated mice were pre-injected intraperitoneally (i.p.) with either AAV-overexpression plasmid (AAV-TBG-Atf4/Ddit3/Btg2) or AAV-sgRNA knockdown plasmid (IP479-EFS-sgBtg2) for 3 weeks. AAV-TBG-EGFP and IP479-EFS-sgEGFP were used as controls, respectively. ALT measurement was performed using a diagnostic assay kit (Teco Diagnostics, Anaheim CA) to assess liver injury.

Added in the revised version: Histology, Immunohistochemistry & immunofluorescence

Added in the revised version: H&E staining

Added in the revised version: Tissues were fixed overnight at 4°C using buffered 10% paraformaldehyde (Sangon Biotech, A500684-0500) and subsequently embedded in paraffin. For H&E staining, paraffin-embedded slides were deparaffinized and rehydrated following standard protocols. The sections were briefly immersed in hematoxylin for 30 seconds, rinsed in water, and then stained with Eosin solution for 3 minutes. After staining, the sections were dehydrated in ethanol, cleared in xylene, and mounted. Slides were examined under a microscope (Scan System, SQS-1000) and analyzed using ImageJ. Necrotic areas outlined by loss of cellular architecture on H&E.

Added in the revised version: Immunofluorescence on frozen section

Added in the revised version: Immunofluorescence staining was conducted on frozen sections of fresh liver tissues. Initially, the tissues were fixed using 2% paraformaldehyde for 1 h and subsequently dehydrated overnight in a 30% sucrose solution. The following day, tissue embedding was performed using OCT (SAKURA, 4583), after which ultra-thin sections of 5μm were sliced. Permeabilization was achieved using 0.2% Triton X-100 for 10 minutes at room temperature, followed by blocking with 5% normal goat serum (VivaCell, C2530-0100). Primary antibodies against mouse GS (Proteintech, 11037-2-AP, 1:200), Chop (Proteintech, 15204-1-AP, 1:1000), Cyp1a2 (Proteintech, 15204-1-AP, 1:50) and Ki67 (Abcam, ab180569, 1:1000) were then applied. Secondary antibodies (Affinipure Goat Anti-Rabbit 488-conjugated, Jackson ImmunoResearch, 111-545-003, 1:400; Goat anti-Rabbit IgG(H+L) Cross Adsorbed Secondary Antibody,Alexa Fluor 568, Invitrogen, A11011, 1:1000 and Goat Anti-Mouse 594-conjugated, Jackson ImmunoResearch, 115-585-003, 1:200) were used accordingly. After washing, the slides were mounted using antifade medium, and nuclei were stained with DAPI (Beyotime, C1006, prediluted). Images were captured using an Olympus BP80 microscope and analyzed with ImageJ.

Added in the revised version: Immunohistochemistry

Added in the revised version: Livers were fixed in 10% buffered formalin for 24 hours before paraffin embedding. For immunohistochemical staining of paraffin-embedded liver tissue, 5 μm sections were prepared, deparaffinized, and subjected to antigen retrieval by microwaving for 20 minutes. Antigen retrieval was performed in pH 6.0 sodium citrate buffer (Ki-67, P62, Chop, Cyp1a2, Cyp2f2), pH 6.4 sodium citrate buffer (p-eIF2A), or pH 9.0 Tris-EDTA buffer (Atf4, Btg2). Sections were blocked with 3% hydrogen peroxide and permeabilized with 0.2% Triton X-100 for 10 minutes at room temperature, followed by blocking with 5% normal goat serum (VivaCell, C2530-0100) for 1 hour at room temperature. Primary antibodies were incubated overnight at 4°C at the following dilutions in blocking buffer: Ki-67 (Abcam, ab15580, 1:400), P62 (Abclonal, A19700, 1:1000), Cyp2f2 (SantaCruz, sc-374540, 1:200), Chop (Proteintech, 15204-1-AP, 1:200), Cyp1a2 (Santa Cruz, sc-53241, 1:50), p-eIF2A (CST, 3398S, 1:50), Atf4 (CST, 11815S, 1:200), and Btg2 (Proteintech, 22339-1-AP, 1:100). Secondary antibodies used were Goat anti-Rabbit IgG (H+L) Secondary Antibody, Biotin (Invitrogen, 65-6140, 1:1000) and peroxidase-AffiniPure Goat Anti-mouse IgG(H+L) (Jackson ImmunoResearch, 115-035-003, 1:1000). Avidin, NeutrAvidin™, Horseradish Peroxidase conjugate (Invitrogen, A2664, 1:2000) was used as the third antibody. After three washes with PBS, color development was achieved using the DAB Peroxidase Substrate Kit (ZSGB-BIO, ZLI-9018), followed by quenching in distilled water. Slides were counterstained with hematoxylin (BBI, A600701-0050), dehydrated to xylene, and images were captured using an Olympus BP80 microscope or Scan System SQS-1000 and analyzed with ImageJ.

Added in the revised version: Western blot analysis

Added in the revised version: Mouse liver tissue (30 mg) was homogenized in 1 mL RIPA buffer (10 mM Tris–HCl pH 8.0, 1 mM EDTA, 0.5 mM EGTA, 1% Triton X-100, 0.1% sodium deoxycholate, 0.1% SDS, 140 mM NaCl, 1 mM PMSF, and protease inhibitor cocktail). The lysate was mixed with 5× SDS loading buffer (final 1×), boiled at 98°C for 10 min, and centrifuged at 12,000 rpm for 3 min at room temperature. Supernatants were separated by 10% SDS-PAGE, transferred to PVDF membranes, and blocked with 5% nonfat milk in TBST (Tris-buffered saline with 0.1% Tween-20). Membranes were incubated with primary antibodies against Cyp2e1 (Proteintech, 19937-1-AP, 1:5000) and GAPDH (Proteintech, 60004-1-Ig, 1:5000), followed by HRP-conjugated secondary antibodies: goat anti-mouse (Jackson ImmunoResearch, 115-035-003, 1:2000) and goat anti-rabbit (Jackson ImmunoResearch, 111-035-003, 1:2000). Protein bands were detected using enhanced chemiluminescence (ECL) and a Minichemi Chemiluminescence Imaging System.

Added in the revised version: RNA extraction and quantitative real-time PCR

Added in the revised version: Ground mouse liver tissue was lysed in TRIzol reagent (Invitrogen, 15596018) and centrifuged at 12,000 rpm for 10 min at 4°C. The supernatant was transferred to a new tube, mixed with 200 μL chloroform, incubated at room temperature for 10 min, and centrifuged. The aqueous phase was collected, mixed with an equal volume of isopropanol, incubated at 4°C for 10 min, and centrifuged. The supernatant was discarded, and the RNA pellet was washed twice with 75% ethanol, air-dried for 10 min, and dissolved in RNase-free water. RNA concentration and purity were measured using a NanoDrop spectrophotometer. Reverse transcription was performed using a cDNA first-strand synthesis kit (Takara, 6210B). Quantitative PCR was carried out in triplicates using PowerUp™ SYBR™ Green Master Mix (Thermo Fisher, A25742) on a QuantStudio 1 Real-Time PCR system (Life Technologies) according to the manufacturer’s instructions. Primer sequences are provided in Table S1.

Added in the revised version: AAV production and purification

Added in the revised version: AAV8 was produced using 293T cells cultured in one or more 175 cm2 cell culture flasks. Cells were plated one day prior to transfection, reaching 50% confluence, which allowed them to reach 80-90% confluence the next day. For transfection of one 175 cm2 flask, a mixture of 11.56 μg pHelper, 7.24 μg pAAV2/8 (Addgene, 112864), and 5.9 μg transgene plasmids was prepared in 3 ml Opti-MEM medium in a tube. PEI solution (1 μg/μl in water, pH 7.0, powder from Polysciences, 23966-2) was added to the tube at a ratio of 3 μg PEI per 1 μg DNA. The solution was mixed and incubated for 15 minutes before adding to the cell culture. At 24 hours post-transfection, the medium was replaced with antibiotic-containing medium (P+/S+) without FBS. At 72 hours after the medium change, cells were scraped off the flask, and 3 ml of chloroform was added, followed by vortexing for 5 minutes. Then, 7.6 ml of 5M NaCl was added and vortexed for 10 seconds. The solution was centrifuged at 3000g for 5 minutes at 4°C. The supernatant was transferred to a new 50 ml tube, and 9.4 ml of 50% (w/v) PEG 8000 was added, vortexed for 10 seconds, and left for 1 hour at 4°C. Subsequently, the mixture was centrifuged at 3000g for 30 minutes at 4°C. The supernatant was decanted, and the tube was inverted for 10 minutes to dry. Next, 1.4 ml of 50 mM HEPES buffer (pH 8.0) was added to re-suspend the pellet, and the solution was vortexed for 5 minutes. Then, 3.5 μl of 1M MgCl2, 2.8 μl of 5 μg/μl DNase I, and 2.8 μl of 5 μg/μl RNase A were added, and the mixture was incubated for 20 minutes at 37°C. The virus was then extracted with chloroform, and the solvent was replaced with PBS. AAV was concentrated by centrifugation at 14000g for 5 minutes at 4°C using a 100kDa ultrafiltration tube (Beyotime, FUF058). Finally, the samples were aliquoted and stored at -80°C after titer measurement.

Added in the revised version: Library preparation for 10x Visium spatial transcriptomics

Added in the revised version: After administering APAP, mice were euthanized at 0h, 3h, 6h and 24h time points using carbon dioxide. Freshly harvested liver tissues were frozen using isopentane and liquid nitrogen, embedded in OCT, sectioned into 10 μm slices, and mounted on Visium tissue optimization or spatial gene expression slides. Tissues were fixed with methanol, stained with H&E, and imaged at 40× magnification. Tissue permeabilization was optimized for 6 minutes. Libraries were prepared according to the Visium spatial gene expression user guide, with cDNA amplification cycles adjusted per time point, and index PCR performed for 6 cycles. Libraries were pooled, loaded at 400 pM, and sequenced on a NovaSeq 6000 SP100 sequencer.

Added in the revised version: Processing of raw sequencing data

Added in the revised version: The raw data underwent processing using spaceranger (version 1.1.0) to perform several analytical tasks. These included detecting tissue boundaries, aligning reads to the reference genome, generating feature-spot matrices, conducting clustering to identify spatially distinct groups of cells, and analyzing gene expression patterns. Additionally, spaceranger facilitated the spatial placement of spots within the context of the slide image, providing a comprehensive spatial transcriptomic analysis of the tissue sample.

Added in the revised version: Basic analysis of spatial transcriptomics (ST) data

Added in the revised version: For the gene-spot matrices generated by Spaceranger, we utilized Scanpy (version 1.9.3) for analysis (Patro et al., 2017). Initially, routine statistical analyses were conducted, including calculating the number of detected UMIs (nUMI) and genes (nGene) in each spot. Basic quality control (QC) measures were then applied to the data. Specifically, spots with extremely low UMI or gene counts (min_counts=2000, min_cell=10) were excluded using scanpy.pp.filter_cells function.

Added in the revised version: Mitochondrial and hemoglobin genes were filtered out to focus on relevant gene expression. Spot matrix was filtered out to keep only spots overlaying tissue sections. Following QC, data normalization was performed using the scanpy.pp.normalize_total function, and log-transformation was applied using the scanpy.pp.log1p function. Subsequently, we identified 2000 highly variable genes using the scanpy.pp.highly_variable_genes function based on their expression means and variances.

Added in the revised version: Principal component analysis (PCA) was then applied to reduce the dimensionality of the data, projecting the spots into a lower-dimensional space using the identified principal components (PCs). Using the corrected PC matrices, we conducted unsupervised clustering based on shared nearest neighbors (SNN) and visualized the data using UMAP (Uniform Manifold Approximation and Projection) for further exploration and analysis.(Lim and Qiu, 2023)

Added in the revised version: Each sample was partitioned into three clusters. To facilitate comparison, each cluster was annotated with a region label (PC, PP, or Mid), determined by integrating information from cluster-specific marker genes and H&E staining images. Furthermore, to compare gene expression profiles across clusters, we identified differentially expressed genes using fold change analysis and the Wilcoxon rank sum test. This approach allowed us to highlight genes showing significant expression differences among all or selected clusters, providing insights into biological variations across spatial regions in the tissue samples.

Added in the revised version: Analysis of differentially expressed genes (DEGs) and pathway enrichment

Added in the revised version: For differential gene analysis across different samples, we initially integrated all samples using the combat function to mitigate batch effects.(Stuart et al., 2019) Subsequently, differential expression analysis of variable genes was conducted using the FindMarkers function, employing the Wilcoxon rank sum test.(Ritchie et al., 2015) Genes showing significant changes were selected based on criteria including a log-fold change threshold of 0.1 and a minimum percentage of samples showing differential expression of 1%. The differentially expressed genes (DEGs) identified were enriched across various time points post-APAP injection, relative to the 0-hour time point.

Added in the revised version: Pathway analysis was performed using the Kyoto Encyclopedia of Genes and Genomes (KEGG), while gene ontology (GO) analysis was carried out using the clusterProfiler package(Yu et al., 2012) to explore functional categories associated with these DEGs. The most significantly enriched pathways were identified using the enrichGO and enrichKEGG functions within the clusterProfiler package.(Yu et al., 2012) Furthermore, Gene Set Variation Analysis (GSVA) was employed to investigate the expression variations related to proliferation at different time points following APAP administration, utilizing the gsva function from the GSVA R package.(Hänzelmann et al., 2013) Pathway activities were visualized using the pheatmap R package for comprehensive analysis and interpretation. To functionally classify the DEGs identified in the F2A comparison, we performed Gene Ontology (GO) enrichment analysis using [name of tool/database, e.g., DAVID, clusterProfiler]. Each DEG was mapped to GO terms from the three ontologies: biological process, cellular component, and molecular function. When a single gene was annotated with multiple GO terms, we selected the most prominent term—defined as the one with the smallest adjusted p- value or highest enrichment score—for downstream interpretation. This approach ensured that the primary functional role of each DEG was captured without redundancy.

Added in the revised version: Gene regulatory network analysis

Added in the revised version: The Single-cell regulatory network inference and clustering (SCENIC) workflow comprises three main steps: coexpression analysis, target gene motif enrichment analysis, and regulon activity assessment. The analysis employed pySCENIC (version 0.12.1)(van de Sande et al., 2020) with default parameters, using the raw count matrix from all samples as input. Co-expression modules were initially identified, and the interaction strength between transcription factors (TFs) and their target genes was evaluated using GRNBoost.(Aibar et al., 2017) Each coexpressed module underwent motif enrichment analysis using RcisTarget with a rank threshold of 3000, to identify modules where the TF motif was significantly enriched among targets.(Aibar et al., 2017) Only modules meeting these criteria were retained, establishing TFs along with their potential direct targets as regulons. The activity of each regulon within each cell was assessed using AUCell. For visualization, average regulon activity scores (AUC) were calculated for each cluster.(Aibar et al., 2017) A rank plot of regulons was generated using ggplot2. Additionally, the gene expression network related to the target genes of the TFs of interest was visualized using the nx.draw function.

Added in the revised version: Analysis of snRNA-seq Data (GSE223561)

Added in the revised version: To annotate cell clusters in the single-nucleus RNA-seq (snRNA-seq) data, we utilized a curated list of known marker genes corresponding to liver cell lineages, as reported in Multimodal Decoding of Human Liver Regeneration. Using the AddModuleScore function in Seurat, we calculated signature scores to assign cell lineage identities. Clusters identified as primarily composed of cycling cells were further reclustered to segregate them into their constituent lineages. This process was repeated iteratively for each identified lineage. A cleansing step was included to remove clusters enriched for nuclei identified as doublets or overexpressing marker genes from multiple lineages. Hepatocytes were subsequently extracted for further reclustering, using resolution (res) = 0.2 and 80 principal components (npc). To annotate the hepatocyte clusters, we applied a curated list of known zonation-specific marker genes for hepatocytes.

Added in the revised version: Library preparation for Cut&Run samples

Added in the revised version: Cut&Run was conducted on mouse liver tissue samples from mice injected with APAP for 6 hours using the Cut&Run assay kit (CST, 86652) following the manufacturer’s protocol. Initially, 1 mg of fresh tissue was weighed for each antibody/MNase reaction, with an additional 1 mg for the input sample. Reactions included positive and negative controls. Tissues were fixed with 1 ml of 0.1% formaldehyde for 10 minutes and quenched with glycine for 5 minutes. After washing with PBS, tissues were resuspended in 1 ml of 1X Wash Buffer (+ spermidine + protease inhibitor cocktail) and transferred to a Dounce homogenizer. Tissue was homogenized into a single-cell suspension with 10-15 strokes until no tissue chunks remained.

Added in the revised version: For the Cut&Run assay to identify Chop binding sites, samples were bound to concanavalin A beads and incubated overnight at 4°C with the following primary antibodies: 0.5 μg CHOP (CST, 2895S, 1:100), 2 μl positive control Tri-Methyl-Histone H3 (Lys4) (CST, 9751), and 5 μl negative control Rabbit (DA1E) mAb IgG XP® Isotype Control (CST, 66362). The pAG MNase enzyme was activated by adding cold calcium chloride and incubating at 4°C for 30 minutes. The reaction was terminated with 1X stop buffer at 37°C for 10 minutes without shaking to release DNA fragments into the solution.

Added in the revised version: To reverse crosslinks in fixed tissue samples, samples were brought to room temperature and treated with 10% SDS Solution (BBI, B548118-0100) to a final concentration of 0.1%, along with proteinase K (20 mg/ml). DNA Extraction Buffer (+ Proteinase K + RNAse A) was added to the input sample. Cells were lysed and chromatin fragmented by sonication using a Covaris ME220 contact ultrasound apparatus on ice, with optimal conditions generating chromatin fragments ranging from 100 to 600 bp. DNA was purified using phenol/chloroform extraction followed by ethanol precipitation, and DNA quantification was performed by qPCR. For Cut&Run, purified DNA was used to prepare sequencing libraries with the NEB Next Ultra II DNA Library Prep Kit for Illumina (7645S and 7335L) and sequenced on an Illumina Nova-Seq 6000 sequencer to obtain 150 bp paired-end reads.

Added in the revised version: Cut&Run data analysis

Added in the revised version: For visualization purposes, BW files were generated using Integrative Genomics Viewer (IGV) software (Robinson et al., 2011), and bigwig files were created using Deeptools (version 3.5.1) with the ‘bamCoverage’ module (bamCoverage --binSize 100) (Ramírez et al., 2016). For each biological replicate and its corresponding IgG control, peaks were called using macs2 (version 2.2.7) with the command macs2 callpeak -f BAMPE --qvalue 0.1 --keep-dup (Zhang et al., 2008). To annotate the identified peaks, ChIPseeker (version 1.26.2) in R was utilized (Yu et al., 2015).

Added in the revised version: Quantification of zonal distribution

Added in the revised version: The zonation of protein (such as Ki67, Chop, Atf4) -positive nuclei was quantified using the following method: the position index (P.I.) was calculated based on distances to the nearest CV (x), portal vein (PV) (y), and the distance between CV and PV (z), utilizing the law of cosines. The formula employed was P.I. = (x^2 + z^2 - y^2) / (2z^2). This approach aligns with the methodology described by Lin et al (Lin et al., 2018).

Added in the revised version: Statistical analysis

Added in the revised version: All experimental data are presented as the mean ± SD. Statistical analyses were performed using GraphPad Prism (v8.0.2), with appropriate tests selected based on experimental design: unpaired two-tailed Student’s t-tests for comparisons between two groups, and one-way ANOVA for comparisons involving three or more groups. The number of animals (‘‘n’’) used in each experiment is indicated in the Figures and corresponding legends. For quantification of liver sections, three to five random pericentral and periportal fields of each liver sample, unless specified otherwise, were imaged and quantified using Image J.

Removed in the revised version: Detailed methods can be found in additional files.